Two-way Dispatched function in Sonic hedgehog shedding and transfer to high-density lipoproteins

Author response:

The following is the authors’ response to the current reviews.

Reviewer #1 (Public Review):

Major shortcomings include the unusual normalization strategies used for many experiments and the lack of quantification/statistical analyses for several experiments. Because of these omissions, it is difficult to conclude that the data justify the conclusions. The significance of the data presented is overstated, as many of the experiments presented confirm/support previously published work. The study provides a modest advance in the understanding of the complex issue of SHH membrane extraction.

Major shortcomings include the unusual normalization strategies used for many experiments and the lack of quantification/statistical analysis for several experiments.

This statement is not correct for the revised manuscript: The normalization strategies used are clearly described in the manuscript and are not unusual. Each experiment is now statistically analyzed.

The significance of the data presented is overstated, as many of the experiments presented confirm/support previously published work.

As reviewer 2 correctly points out, there are many competing models for Hedgehog release. Our study cannot possibly support them all - the reviewer's statement is therefore misleading. In fact, our careful biochemical analysis of the mechanistics of Dispatched- mediated Shh export supports only two of them: The model of proteolytic processing of Shh lipid anchors (shedding) and the model of lipoprotein-mediated Shh transport. In contrast, our study does not support the predominant model of Dispatched-mediated extraction of dual-lipidated Shh and delivery to Scube2, which is currently thought to act as a soluble Shh chaperone. We also do not support Dispatched function in Shh endocytic recycling and cytoneme loading, or any of the other models such as exosome-mediated or micelle Shh transport.

Reviewer #2 (Public Review):

A novel and surprising finding of the present study is the differential removal of Shh N- or C- terminal lipid anchors depending on the presence of HDL and/or Disp. In particular, the identification of a non-palmitoylated but cholesterol-modified Shh variant that associates with lipoproteins is potentially important. The authors use RP-HPLC and defined controls to assess the properties of processed forms of Shh, but their precise molecular identity remains to be defined. One caveat is the heavy reliance on overexpression of Shh in a single cell line. The authors detect Shh variants that are released independently of Disp and Scube2 in secretion assays, but these are excluded from interpretation as experimental artifacts. Therefore, it would be important to demonstrate key findings in cells that endogenously secrete Shh.

We would like to respond as follows:

The authors use RP-HPLC and defined controls to assess the properties of processed forms of Shh, but their precise molecular identity remains to be defined.

This is the original reviewers statement regarding our original manuscript submission. We believe that the biochemical and functional data presented in the VOR clearly describe the molecular identity of solubilized Shh: it is monolipidated, lipoprotein-associated, and highly biologically active in two established Shh bioassays.

One caveat is the heavy reliance on overexpression of Shh in a single cell line.

As stated by reviewer 1, the strength of our work is the use of a bicistronic SHH-Hhat system to consistently generate doubly lipidated ligand to determine the amount and lipidation status of SHH released into cell culture media. This unique system therefore eliminates the artifacts of protein overexpression. We have also added two other cell lines to our VOR that produce the same results (including Panc1 cells that endogenously produce Shh, Supplementary Figure 1).

The authors detect Shh variants that are released independently of Disp and Scube2 in secretion assays, but these are excluded from interpretation as experimental artifacts.

As the reviewer correctly points out, these variants are released independently of Disp and Scube2, both of which are known as essential release factors in vivo. These variants are therefore by definition experimental artifacts. The forms we have included in our analysis are the alternative forms that are clearly dependent on Dispatched and Scube2 for their release - as shown in the first figure in the manuscript, and in pretty much every other figure after that.

The following is the authors’ response to the previous reviews.

Reviewer #1 (Public Review):

Key shortcomings include the unusual normalization strategies used for many experiments and the lack of quantification/statistical analyses for several experiments.

In the updated version of the paper, we have addressed all of this reviewer's criticisms. Most importantly, we have performed several additional experiments to address the concern that unusual normalization strategies were used in our paper and that quantification and statistical analyses were lacking for several experiments. We have now analyzed the full set of release conditions for Shh and engineered proteins from Disp-expressing n.t. control cells and Disp-/- cells both in the presence and absence of Scube2 (Figure 1A'-D', Figure 2E added to the paper, Figure 3B'-D', Figure 5C and Figure S2F-H). Previously, we had only quantified protein release from n.t. controls and Disp-/- cells in the presence but not in the absence of Scube2 under serum-depleted conditions. Quantifications of serum-free protein release and Shh release under conditions ranging from 0.05% FCS to 10% FCS were completely missing from the earlier versions of the manuscript, but have now been added to our paper. In addition, we have reanalyzed all of the data sets in the above figures, as well as Figures 2C and S1B, to address the issue of "unusual normalization strategies": unlike previous assays in which the highest amount of protein detected in the media was set to 100% and all other proteins in that experiment were expressed relative to that value, we now directly compare the relative amounts of cellular and corresponding solubilized proteins as a method to quantify release without the need for data normalization (Figs. 1A'-D', 2C,E, 3B'-D', E, 5C, Fig. S1B, S2F-H).

We have also repeated the qPCR analyses in C3H10T1/2 cells and now show that the same Shh/C25AShh activities can be observed when using another Shh responsive cell line, NIH3T3 cells (Fig. 4B, 6B, fig. S5B).

We would like to point out that if the criticism refers to the presentation of our RP-HPLC and SEC data, the normalization of the strongest eluted protein signal to 100% for all proteins tested is necessary to put their behavior in a clearer relationship. This is because only the relative positions of protein elution, and not their amounts, are important in these experiments.

The significance of the data provided is overstated because many of the presented experiments confirm/support previously published work.

To mitigate the first reviewer's comment that the significance of the data presented is overstated, we now clearly distinguish between our novel results and the known aspect of Hh release on lipoproteins throughout our paper. We now clearly describe what is new and important in our paper: First, contrary to the general perception in the field, Disp and Scube2 are not sufficient to solubilize Shh, casting doubt on the currently accepted model that Scube2 accepts dual-lipidated Shh from Disp and transports it to the receptor Ptch. Second, lipoproteins shift dual Shh processing to N-terminal peptide processing only to generate different soluble Hh forms with different activities (as shown in Figure 4C). Third, and again contrary to popular belief, this new release mode does not inactivate Shh, as we now show in two established cellular assays for Hh biofunction (Figures 4A-C, 5B'', 6B and S5C-G). Fourth, and most importantly, we show that spatiotemporally controlled, Disp-, Scube2- and HDL-mediated Shh release absolutely requires dual lipidation of the membrane-associated Shh precursor prior to its release. This finding (as shown in Figures 1 and S2) changes the interpretation of previously published in vivo data that have long been interpreted as evidence for the requirement of dual Shh lipidation for full receptor binding and activation.

The study provides a modest advance in our understanding of the complex issue of Shh membrane extraction.

Although we agree that our results integrate our novel observations into previously established concepts of Hh release and trafficking, we also hope that our data cast well-founded doubt on the current view that the issue of Hh release and trafficking is largely resolved by the model of Disp-mediated Shh hand-over to Scube2 and then to Ptch, which requires interactions with both Shh lipids. Our data show that this is clearly not the case in the presence of lipoproteins. Thus, the significance of our data is that models of Shh lipid-regulated signaling to Ptch obtained using the dual-lipidated Shh precursor prior to its Disp- and Scube2-mediated conversion into a delipidated or monolipidated, HDL-associated soluble ligand are likely to describe a non-physiological interaction. Instead, our work describes a highly bioactive soluble ligand with only one lipid still attached, which has not been described before in the literature. The in vivo endpoint analyses presented in Fig. S8 suggest that this new protein variant is likely to play an important role during development.

Reviewer #2 (Public Review):

The precise molecular identity (of the released Shh) remains to be defined.

We would like to respond that the direct comparison of soluble proteins and their well-defined double-lipidated precursors side-by-side in the same experiment, as shown in our paper, determines all relevant molecular changes in the Shh release process. Most importantly, we show by SDS-PAGE and RP-HPLC that HDL restricts Shh processing to the N-terminus and that the absence of HDL results in double processing of Shh during its release. We also show by SEC that the C-terminus binds the protein to HDL. In addition, the fly experiments confirm the requirement for N-terminal Hh processing, but not for processing of the C-terminal peptide, and suggest that the N-terminal Cardin-Weintraub sequence replaced by the functionally blocking tag represents the physiological cleavage site.

It would be important to demonstrate key findings in cells that secrete Shh endogenously.

We now confirm the key findings of our study in Panc1 cells that endogenously produce and secrete Shh: As shown in Fig. S1D, we find that soluble proteins are processed but retain the C-cholesterol, which we now directly confirm by RP-HPLC (Fig. S4F-H). The in vivo analyses shown in Fig. S8 suggest that the key finding - that N-terminal but not C-terminal Hh shedding is required for release - can be supported, at least in the fly: here, Hh variants impaired in their ability to be processed N-terminally strongly repress the endogenous protein, whereas the same protein impaired in its ability to be processed C-terminally does not.

The authors detect Shh variants that are expressed independently of Disp and Scube2 in secretion assays, but are excluded from interpretation as experimental artifacts.

We agree with the reviewer's criticism that the amounts of Shh released independently of Disp and Scube2 in secretion assays were not quantified and analyzed statistically to justify their proposed status as not physiologically relevant. We now show that these forms are indeed secretion artifacts (Fig. 3E and Fig. S2F-H show quantification of the lower electrophoretic mobility protein fraction (i.e., the "top" band representing the double-lipidated soluble protein fraction)) because this fraction is released independently of Disp and Scube2.

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Summary

Earlier version: The Sonic hedgehog (Shh) signaling pathway controls embryonic development and tissue homeostasis after birth. This requires regulated solubilization of dual-lipidated, firmly plasma membrane-associated Shh precursors from producing cells. Although it is firmly established that the resistance-nodulation-division transporter Dispatched (Disp) drives this process, it is less clear how lipidated Shh solubilization from the plasma membrane is achieved. We previously showed that Disp enhances proteolytic Shh solubilization from its lipidated terminal peptide anchors. This process, called shedding, converts tightly membrane-associated hydrophobic Shh precursors into delipidated soluble proteins. We show here that Disp-mediated Shh shedding is modulated by a serum factor that we identify as high-density lipoprotein (HDL). In addition to serving as soluble sinks for free membrane cholesterol, HDLs also accept the cholesterolmodified Shh peptide from Disp. The cholesteroylated Shh peptide is required and sufficient for Disp-mediated transfer because mCherry linked to cholesteroylated peptides associates with HDL in a Disp-dependent manner, but an N-palmitoylated Shh variant that lacks C-cholesterol does not. Disp-mediated Shh transfer to HDL is finalized by proteolytic processing of the palmitoylated N-terminal membrane anchor. Unlike dual processed soluble Shh with moderate bioactivity, HDL-associated N-processed Shh is highly bioactive. We suggest that the purpose of different post-translationally modified soluble Shh forms generated from one dual-lipidated cellular precursor is to fine-tune cellular responses in a tissue type-specific and time-specific manner.

Revised version: The Sonic hedgehog (Shh) signaling pathway controls embryonic development and tissue homeostasis after birth. This requires regulated solubilization of dual-lipidated, firmly plasma membrane-associated Shh precursors from producing cells. Although it is firmly established that the resistance-nodulation-division transporter Dispatched (Disp) drives this process, it is less clear how lipidated Shh solubilization from the plasma membrane is achieved. We have previously shown that Disp promotes proteolytic solubilization of Shh from its lipidated terminal peptide anchors. This process, termed shedding, converts tightly membrane-associated hydrophobic Shh precursors into delipidated soluble proteins. We show here that Disp-mediated Shh shedding is modulated by a serum factor that we identify as high-density lipoprotein (HDL). In addition to serving as a soluble sink for free membrane cholesterol, HDLs also accept the cholesterol-modified Shh peptide from Disp. The cholesteroylated Shh peptide is necessary and sufficient for Disp-mediated transfer because mCherry linked to cholesteroylated peptides associates with HDL in a Disp-dependent manner, whereas an N-palmitoylated Shh variant lacking C-cholesterol does not. Disp-mediated Shh transfer to HDL is completed by proteolytic processing of the palmitoylated N-terminal membrane anchor. In contrast to dual processed soluble Shh with moderate bioactivity, HDL-associated N-processed Shh is highly bioactive. We suggest that the purpose of different post-translationally modified soluble Shh forms generated from the dual-lipidated cellular precursor is to fine-tune cellular responses in a tissue type and time-specific manner.

Introduction

Earlier version: Hedgehog (Hh) ligands activate an evolutionarily conserved signaling pathway that provides instructional cues during tissue morphogenesis, which, if misregulated, contributes to developmental disorders and cancer. Hhs are unique in that they require autocatalytic covalent cholesteroylation of the C-terminus 1 and N-terminal palmitoylation by a separate Hh acyltransferase (Hhat) activity 2. Both lipids firmly tether Hh to the plasma membrane of producing cells to effectively prevent unregulated ligand release 3. Signaling at distant target cells therefore requires the regulated solubilization of Hh from the membrane of ligand-producing cells, a process that is facilitated by vertebrate and invertebrate Dispatched (Disp) orthologs in vitro 4 and in vivo 5, 6, 7, 8. Disp contains 12 transmembrane helices and 2 extracellular domains and belongs to the resistance-nodulation-division family of transmembrane efflux pumps. Normally, such pumps maintain cellular homeostasis and remove toxic compounds. In addition, Disp contains a conserved multipass transmembrane domain known as the sterol-sensing domain (SSD) that regulates the homeostasis of free or esterified cellular cholesterol in other SSD proteins 9. This molecular architecture is in line with Disp extraction of free cholesterol from the plasma membrane to remove it from the cell 4. In addition, given its established role in Hh/Shh release, Disp has also been suggested to extract the C-terminal Shh sterol to solubilize Hh/Shh into the extracellular compartment 10.

Revised version: Hedgehog (Hh) ligands activate an evolutionarily conserved signaling pathway that provides instructional cues during tissue morphogenesis and, when misregulated, contributes to developmental disorders and cancer. Hhs are unique in that they require autocatalytic covalent cholesteroylation of the C-terminus 1 and N-terminal palmitoylation by a separate Hh acyltransferase (Hhat) activity 2. Both lipids tightly bind Hh to the plasma membrane of producing cells to effectively prevent unregulated ligand release 3. Signaling to distant target cells therefore requires the regulated solubilization of Hh from the membrane of ligand-producing cells, a process that is facilitated by vertebrate and invertebrate Dispatched (Disp) orthologs in vitro 4 and in vivo 5, 6, 7, 8. Disp contains 12 transmembrane helices and 2 extracellular domains and belongs to the resistance-nodulation-division family of transmembrane efflux pumps. Normally, such pumps maintain cellular homeostasis and remove toxic compounds. In addition, Disp contains a conserved multi-pass transmembrane domain known as the sterol sensing domain (SSD), which regulates the homeostasis of free or esterified cellular cholesterol in other SSD proteins 9. This molecular architecture is consistent with Disp extraction of free cholesterol from the plasma membrane to remove it from the cell 4. In addition, given its established role in Hh/Shh release, it has been suggested that Disp also extracts the C-terminal Shh sterol to solubilize Hh/Shh in the extracellular compartment 10.

Earlier version: However, Disp activity alone is insufficient to solubilize the vertebrate Hh family member Shh from the plasma membrane. A second required synergistic factor for maximum Shh signaling is the soluble extracellular glycoprotein Scube2 (Signal sequence, cubulin (CUB) domain, epidermal growth factor (EGF)-like protein 2) 9. One way to explain Disp/Scube2 synergy is by Disp-mediated extraction of dual-lipidated Shh from the plasma membrane to hand it over to Scube2, a mechanism that depends on the C-terminal cysteine-rich and CUB domains of Scube2 10, 11. In vitro support for this mechanism comes from Disp and Scube2 co-immunoprecipitation with Shh 10, 11 and from structural data 12. Other suggested carriers for Dispextracted lipidated Hh/Shh include lipoproteins (LPPs) 13, 14, 15, exosomes as carriers of internalized and re-secreted Hhs 16, 17, 18, and micellous Hh assemblies 19. Another suggested Shh release mode is Disp-regulated proteolytic processing (called shedding) from the plasma membrane 4 by two major plasma-membrane-associated sheddases called A Disintegrin and Metalloproteinase (ADAM) 10 and 17 20, 21, 22. The role of Scube2 in the shedding model is to enhance proteolytic processing of both terminal lipidated Shh peptides in a CUB domaindependent manner 23, 24, 25. Consistent with this, CUB function in regulated shedding of other substrates often involves substrate recognition and induced structural changes in the substrate to boost turnover 25, 26, 27.

Revised version: However, Disp activity alone is not sufficient to solubilize the vertebrate Hh family member Sonic Hedgehog (Shh) from the plasma membrane. A second synergistic factor required for maximal Shh signaling is the soluble extracellular glycoprotein Scube2 (Signal sequence, cubulin (CUB) domain, epidermal growth factor (EGF)-like protein 2) 9. One way to explain the Disp/Scube2 synergy is that Disp-mediated extraction of dual-lipidated Shh from the plasma membrane delivers it to Scube2, a mechanism that depends on the C-terminal cysteine-rich and CUB domains of Scube2 10, 11. In vitro support for this mechanism comes from co-immunoprecipitation of Disp and Scube2 with Shh 10, 11 and from structural data 12. Other proposed carriers for Disp-extracted lipidated Hh/Shh include lipoprotein particles (LPPs) 13, 14, 15, exosomes as carriers of internalized and re-secreted Hhs 16, 17, 18, and micellous Hh assemblies19. Another proposed mode of Shh release is Disp-regulated proteolytic processing (termed shedding) from the plasma membrane 4 by two major plasma membrane-associated sheddases, A Disintegrin and Metalloproteinase (ADAM) 10 and 17 20, 21, 22. The role of Scube2 in the shedding model is to enhance proteolytic processing of both terminal lipidated Shh peptides in a CUB domain-dependent manner 23, 24, 25. Consistent with this, CUB function in regulated shedding of other substrates often involves substrate recognition and induced structural changes in the substrate to increase turnover 25, 26, 27.

Earlier version: In this study, we systematically characterized the biochemical parameters of Dispregulated Shh solubilization with all these published models and parameters in mind. To this end, we used a unique bicistronic Hhat/Shh co-expression system to ensure that only dual-lipidated Shh as a substrate for Disp was produced and analyzed 23. We further avoided any protein tags to ensure undisturbed Shh modification, secretion, and interaction with Disp and other potential pathway components. Finally, we used unbiased biochemical methods to analyze Shh solubilization, size, and lipidation status under different experimental conditions. First, we confirmed by SDS-PAGE/immunoblotting and reverse-phase (RP)-HPLC that Disp and Scube2 synergistically enhance Shh shedding into the cell culture medium 23. We also found that repeated careful washing of Disp- and Shh-expressing cells to remove all serum traces abolished Shh shedding, and that high serum amounts promoted the solubilization of a previously unknown, highly bioactive Shh variant that lacks its palmitoylated N-terminus but retains its cholesteroylated C-terminus. Size exclusion chromatography (SEC) analyses demonstrated co-elution of this novel Shh variant with serum apolipoprotein A1 (ApoA1), the major protein component of the high-density lipoprotein (HDL) fraction in serum. Consistent with this observation, purified HDL restored solubilization of the N-processed, bioactive Shh variant from washed cells. We also determined that the most C-terminal cholesteroylated Shh peptide is sufficient for Disp-mediated protein transfer to HDL. In contrast, palmitoylated N-terminal Shh membrane anchors are not transferred by Disp, but undergo shedding to complete Shh transfer to HDL. Our updated Hh shedding model therefore unifies previously disparate models of Disp-, sheddase-, and LPP-mediated Shh solubilization into one comprehensive system. This comprehensive system is fully in line with published reports of in vivo Disp functions 5, 6, 7, 8, 28, in vivo Scube2 functions 29, 30, 31, in vitro roles of the Scube2 CUB domain 23, 25, the Disp structure 32, 33, and required N-terminal – but not C-terminal – shedding during Hh solubilization in vivo 34, 35, 36. Our model is also in line with the established importance of Hh C-cholesterol for Hh association into “large punctate” structures that are visible by light microscopy and may represent Hh-LPP complexes 37, 38, 39, C-cholesterol-dependent Hh spread 10, 40, and previously established in vivo roles of LPPs in Hh biofunction 13, 14, 15. We suggest that the Disp-produced, HDL-linked mono-lipidated Shh variant presented in this work helps to meet the demands for Hh activity regulation in specific cell types and developing tissues in vivo.

Revised version: In this study, we systematically characterized the biochemical parameters of Disp-regulated Shh solubilization with all these published models and parameters in mind. To this end, we used a unique bicistronic Hhat/Shh coexpression system to ensure that only dual-lipidated Shh was produced and analyzed as a substrate for Disp 23. We also avoided any protein tagging to ensure undisturbed Shh modification, secretion, and interaction with Disp and other potential pathway components. Finally, we used unbiased biochemical methods to analyze Shh solubilization, size, and lipidation status under different experimental conditions. First, we confirmed by SDS-PAGE/immunoblotting and reverse phase (RP)-HPLC that Disp and Scube2 synergistically enhance Shh shedding into the cell culture medium 23. We also found that repeated washing of Disp- and Shh-expressing cells to remove all traces of serum abolished Shh shedding, and that high levels of serum promoted the solubilization of a previously unknown, highly bioactive Shh variant that lacks its palmitoylated N-terminus but retains its cholesteroylated C-terminus. Size exclusion chromatography (SEC) analyses revealed co-elution of this novel Shh variant with serum apolipoprotein A1 (ApoA1), the major protein component of the serum high-density lipoprotein (HDL) fraction. Consistent with this observation, purified HDL restored solubilization of the N-processed, bioactive Shh variant from washed cells. We also found that the most C-terminal cholesteroylated Shh peptide is sufficient for Disp-mediated protein transfer to HDL. In contrast, palmitoylated N-terminal Shh membrane anchors are not transferred by Disp, but undergo shedding to complete Shh transfer to HDL. Our updated Hh shedding model thus unifies previously disparate models of Disp-, sheddase-, and LPP-mediated Shh solubilization into a comprehensive system. This comprehensive system is fully consistent with published reports on in vivo Disp functions 5, 6, 7, 8, 28, in vivo Scube2 functions 29, 30, 31, the in vitro role of the Scube2 CUB domain 23, 25, the Disp structure 32, 33, and the required N-terminal – but not C-terminal – shedding during in vivo Hh solubilization 34, 35, 36. Our model is also consistent with the established importance of Hh C-cholesterol for Hh association into “large punctate” structures visible by light microscopy that may represent Hh-LPP complexes 37, 38, 39, C-cholesterol-dependent Hh spreading 10, 40, and previously established in vivo roles of LPPs in Hh biofunction 13, 14, 15. We propose that the Disp-produced, HDL-bound, mono-lipidated Shh variant presented in this work helps to meet the requirements for regulating Hh activity in specific cell types and developing tissues in vivo.

Results

Earlier version: Synergistic Disp and Scube2 function increases shedding of dual-lipidated Shh from the plasma membrane

Revised version: Synergistic function of Disp and Scube2 increases the shedding of dual-lipidated Shh from the plasma membrane

Earlier version: To analyze Shh solubilization from the plasma membrane, we produced dual-lipidated, tightly plasma membrane-associated morphogens in Bosc23 control cells that express Disp endogenously 23 (called CRISPR non-targeting control (nt ctrl) cells in this study) and in Bosc23 cells made deficient in Disp function by CRISPR/Cas9 (using previously characterized Disp-/- cells 4; experimental outlines are described in Fig. S1). Shh biosynthesis in both cell lines starts with the covalent attachment of a cholesterol moiety to the C-terminus of the Shh signaling domain 41. This reaction is autocatalytic and strictly coupled to the generation of 19 kDa Shh signaling domains from a 45 kDa precursor. Subsequent N-palmitoylation of cholesteroylated Shh/Hh requires a separate enzymatic activity encoded by Hhat 42. Because Bosc23 cells do not express endogenous Hhat 23, we minimized the production of non-palmitoylated Shh by using bicistronic mRNA that encodes the Shh precursor together with Hhat. Unlike Shh expression in the absence of Hhat, coupled Shh/Hhat expression ensures near-quantitative Shh N-palmitoylation in transfected Bosc23 cells 23. SDS-PAGE/immunoblotting was then used to characterize dual-lipidated Shh extraction from the plasma membrane of both cell lines into serum-depleted media (Fig. 1A; Fig. S2 includes all transfection and loading controls). We confirmed the published finding that Scube2 enhances Shh solubilization from nt ctrl cells (arrowhead) and that Shh solubilization from Disp-/- cells was strongly reduced 4, 10, 11 (Fig. 1A’, Table S1). We also confirmed that the electrophoretic mobility of most Shh released from nt ctrl cells was increased over that of the corresponding dual-lipidated cellular material on the same blot 4, 23, 24 (asterisk). RP-HPLC of the solubilized material shown in Fig. 1A demonstrated that the observed increase in electrophoretic mobility was caused by the proteolytic removal of both lipidated terminal peptides (this post-translational Shh modification is called shedding throughout this manuscript, Fig. 1A’’; see Fig. S2 for cellular dual-lipidated or artificially produced monolipidated and unlipidated Shh standard proteins and see reference 4, which ruled out alternative Shh deacylation modes). Transgenic Disp expression in Disp-/- cells restored Shh shedding and solubilization, confirming the specificity of the assay (Fig. S2E) 4. Strikingly, the same phenotypic reversal was achieved by the co-expression of the cholesterol pump Patched (Ptch1) 43, which depletes the plasma membrane of free sterols 4. This suggests an indirect “second messenger” role of plasma membrane cholesterol not only in the regulation of Smoothened downstream of Ptch1 43, 44, 45, but also in the regulation of Shh shedding by Disp (Fig. S2E) 4. Note that alternative solubilized Shh variants (represented by the “upper” bands) are also present in the media. Release of these variants, however, is independent of Disp and Scube2. In contrast, Scube2 and Disp synergistically and specifically increased Shh shedding from the cell surface (as indicated by the appearance of “lower” sized bands compared with those of cellular proteins) 4.

Revised version: To analyze Shh solubilization from the plasma membrane, we produced dual-lipidated, tightly plasma membrane-associated morphogens in Bosc23 control cells that endogenously express Disp 23 (referred to in this study as CRISPR non-targeting control (nt ctrl) cells) and in Bosc23 cells made deficient in Disp function by CRISPR/Cas9 (using previously characterized Disp-/- cells 4; experimental outlines and specificity controls are described in Fig. S1A-E). Shh biosynthesis in both cell lines begins with the covalent attachment of a cholesterol moiety to the C-terminus of the Shh signaling domain 41. This reaction is autocatalytic and tightly coupled to the generation of 19 kDa Shh signaling domains from a 45 kDa precursor. However, the subsequent N-palmitoylation of cholesteroylated Shh/Hh requires a separate enzymatic activity encoded by Hhat 42. Since Bosc23 cells do not express endogenous Hhat 23, we minimized the production of non-palmitoylated Shh by using bicistronic mRNA encoding the Shh precursor together with Hhat. In contrast to Shh expression in the absence of Hhat, coupled Shh/Hhat expression ensures nearly quantitative Shh N-palmitoylation in transfected Bosc23 cells 23. SDS-PAGE/immunoblotting was then used to characterize dual-lipidated Shh extraction from the plasma membrane of both cell lines into serum-depleted media (Fig. 1A; Fig. S2 includes all transfection and loading controls). We confirmed the published finding that Scube2 enhanced Shh solubilization from nt ctrl cells (arrowhead) and that Shh solubilization from Disp-/- cells was strongly reduced 4, 10, 11 (Fig. 1A’, Table S1). We also confirmed that the electrophoretic mobility of most Shh released from nt ctrl cells was increased over that of the corresponding dual-lipidated cellular material on the same blot 4, 23, 24 (Fig. 1A, asterisk). RP-HPLC of the solubilized material shown in Fig. 1A demonstrated that the observed increase in electrophoretic mobility was caused by the proteolytic removal of both lipidated terminal peptides (this post-translational modification of Shh is referred to as shedding throughout this manuscript, Fig. 1A’’; see Fig. S2 for cellular dual-lipidated or artificially produced monolipidated and unlipidated Shh standard proteins and see Reference 4, which rules out alternative modes of Shh deacylation). Transgenic Disp expression in Disp-/- cells restored Shh shedding and solubilization, confirming the specificity of the assay (Fig. S2E) 4. Strikingly, the same phenotypic reversal was achieved by the coexpression of the cholesterol pump Patched (Ptch1) 43, which depletes the plasma membrane of free sterols 4. This suggests an indirect “second messenger” role of plasma membrane cholesterol not only in the regulation of Smoothened downstream of Ptch1 43, 44, 45, but also in the regulation of Shh shedding by Disp (Fig. S2E) 4. Note that alternative solubilized Shh variants (represented by the “top” bands) are also present in the media. However, the release of these variants is independent of Disp and Scube2 (Fig. S2F). In contrast, Scube2 and Disp synergistically and specifically increased Shh shedding from the cell surface (as indicated by the appearance of “lower” sized bands compared to those of cellular proteins, Fig. 1A, A’) 4.

Earlier version: Disp and Scube2 enhance shedding of cell surface-associated dual-lipidated Shh into delipidated soluble forms.

Revised version: Disp and Scube2 enhance the shedding of cell surface-associated dual-lipidated Shh into delipidated soluble forms.

Earlier version: Media containing 10% serum were exchanged for serum-free DMEM 36 h post-transfection and proteins solubilized for 6 h. Cells were not washed between media changes to leave residual serum traces in the assay. We refer to this experimental condition as “serum-depleted” throughout this work. A) Cap-dependent Shh translation and cap-independent Hhat translation from one bicistronic mRNA ensured the generation of dual-lipidated, plasma membrane-associated proteins (asterisk in the cellular (c) fraction) in all transfected cells. Disp and Scube2 synergistically and specifically enhance dual-lipidated Shh precursor conversion into truncated soluble variants during release (arrowhead). m: media, c: cell lysate. A’) Quantification of relative Shh release from nt ctrl and Disp-/- cells in the presence of Scube2. Ratios of solubilized versus cellular Shh were determined and expressed relative to Shh solubilization from nt ctrl cells (set to 100%, black bar). A’’) Reverse-phase HPLC analyses revealed that Shh solubilized by Disp and Scube2 (the same fraction indicated by the arrowhead in A, black line) was less hydrophobic than its cell surface-associated precursor (gray line; the asterisk denotes analysis of the same cellular fraction as shown in A). RP-HPLC calibration and color coding of Shh fractions is shown in Fig. S2 F-K. B-D) Solubilization of non-palmitoylated C25SShh (in this artificial variant, the N-terminal palmitate acceptor cysteine is changed for a non-accepting serine; functionally equivalent constructs with the cysteine exchanged for a non-accepting alanine (C25AShh) were also used in our study), non-cholesteroylated yet palmitoylated ShhN and lipid-free control C25SShhN under the same serum-depleted conditions. Arrowheads denote Shh variants that were solubilized by Disp- and Scube2. B’-D’) Quantifications of B-D, again from nt ctrl and Disp-/- cells in the presence of Scube2. B’’-D’’) Reverse-phase HPLC demonstrates similar elution of C25SShh, ShhN, and non-lipidated C25SShhN. This demonstrates that terminal lipids were removed during C25SShh and ShhN (as well as Shh, A’’) release. Unpaired t-test, two-tailed. ****: p<0.0001, *: p<0.016, n.t.: 0.79. n=20 datasets (A’), n=13 datasets (B’), n=10 datasets (C’), and n=17 datasets (D’). See Table S1 for detailed statistical information.

Revised version: Media containing 10% serum was changed to serum-free DMEM 36 h post-transfection and proteins were solubilized for 6 h. Cells were not washed between media changes to leave residual serum traces in the assay. We refer to this experimental condition as “serum-depleted” throughout this paper. A) Cap-dependent Shh translation and cap-independent Hhat translation from bicistronic mRNA ensured the generation of dual-lipidated, plasma membrane-associated proteins (asterisk in the cellular fraction) in all transfected cells. Disp and Scube2 synergistically and specifically enhance the conversion of dual-lipidated Shh precursors into truncated soluble variants during release (arrowhead). m: media, c: cell lysate. A’) Quantification of relative Shh release from nt ctrl and Disp-/- cells in the presence or absence of Scube2. Amounts of solubilized Shh with higher electrophoretic mobility (the lower bands) were quantified and expressed as % relative to the respective cellular Shh, which was always set to 100%. One-way ANOVA, Dunnett’s multiple comparisons test. ****: p<0.0001. See Table S1 for detailed statistical information. A’’) Reverse phase HPLC analyses revealed that Shh solubilized by Disp and Scube2 (the same fraction indicated by the arrowhead in A, black line) was less hydrophobic than its cell surface-associated precursor (gray line; the asterisk indicates analysis of the same cellular fraction as shown in A). RP-HPLC calibration and color coding of Shh fractions are shown in Fig. S2 I-N. B-D) Solubilization of non-palmitoylated C25SShh (in this artificial variant, the N-terminal palmitate acceptor cysteine is replaced by a non-accepting serine; functionally equivalent constructs with the cysteine exchanged for a non-accepting alanine (C25AShh) were also used in our study), non-cholesteroylated but palmitoylated ShhN and lipid-free control C25SShhN under the same serum-depleted conditions. Arrowheads indicate Shh variants that were solubilized in Disp- and Scube2 presence. B’-D’) Processed protein quantifications from B-D, again from nt ctrl and Disp-/- cells in the presence or absence of Scube2. One-way ANOVA, Dunnett’s multiple comparisons test. ns: p>0.05. See Table S1 for detailed statistical information. B’’-D’’) Reverse phase HPLC shows similar elution of C25SShh, ShhN, and non-lipidated C25SShhN. This indicates that terminal lipids were removed during the release of C25SShh and ShhN (as well as Shh, A’’).

Earlier version: We also observed that solubilization of artificially produced monolipidated Shh variants (C25SShh that lacks N-palmitate, Fig. 1B, B’, and ShhN that lacks C-cholesterol, Fig. 1C, C’, Fig. S2) was much less controlled, consistent with the observations of others 10, 11. Release of both monolipidated Shh variants remained linked to shedding of the respective lipidated membrane termini, as shown by electrophoretic mobility shifts of soluble proteins (Fig. 1B,C, arrowheads) and their delipidation during solubilization (as determined by RP-HPLC, Fig. 1B’’,C’’). Nonlipidated C25SShhN was always secreted in a Disp- and Scube2-independent manner, as expected (Fig. 1D-D’’). These results suggest that cell-surface shedding represents a “ground-state” from which dual-lipidated Shh is protected, but only until Disp and Scube2 render it prone to shedding. They also suggest that the complete conservation of dual N- and C-terminal Hh lipidation during biosynthesis serves to prevent unregulated protein release from producing cells. This demonstrates the importance of coupled Shh/Hhat expression in vitro to reliably characterize the mechanism of Disp- and Scube2-regulated Shh solubilization.

Revised version: We also observed that solubilization of engineered monolipidated Shh variants (C25SShh, lacking N-palmitate, Fig. 1B, B’, and ShhN, lacking C-cholesterol, Fig. 1C, C’, Fig. S2B,C) was much less controlled, consistent with the observations of others 10, 11. The release of both monolipidated Shh variants remained linked with the shedding of the respective lipidated membrane termini, as shown by electrophoretic mobility shifts of soluble proteins (Fig. 1B,C, arrowheads) and their delipidation during solubilization (as determined by RP-HPLC, Fig. 1B’’,C’’). The unprocessed protein fraction, represented by the “top” bands on the immunoblots, was also released in a manner that was independent of Disp and Scube2 (Fig. S2G,H), as was the non-lipidated C25SShhN (Fig. 1D-D’’, Fig. S2D). Taken together, these results suggest that cell surface shedding represents a “ground state” from which dual-lipidated Shh is protected, but only until Disp and Scube2 render it susceptible to shedding. Our results also suggest that the complete conservation of dual N- and C-terminal Hh lipidation during biosynthesis serves to prevent unregulated protein release from producing cells. This highlights the importance of coupled Shh/Hhat expression in vitro to reliably characterize the mechanism of Disp- and Scube2-regulated Shh solubilization.

Earlier version: Shh shedding depends on cleavage-activated Disp and the presence of serum

Revised version: Shh shedding is dependent on cleavage-activated Disp and the presence of serum

Earlier version: It was recently shown that the prohormone convertase furin cleaves Disp at a conserved processing site to activate it and to release Shh from the cell surface 46. On the basis of this activation mode, we hypothesized that furin inhibition may specifically impair Disp-regulated Shh shedding. To test this hypothesis, we added 0-80 μM peptidyl chloromethylketone (CMK, a competitive inhibitor of furin) to our solubilization assays. Indeed, CMK reduced Shh shedding from the cell surface in a concentration-dependent manner (Fig. 2A,B, Fig. S3A). During these assays, we also noted that repeated careful washing of cells to quantitively remove all serum traces strongly impaired Disp- and Scube2-mediated Shh solubilization into serum-free media (Fig. 2C, Fig. S3B). We derive two important conclusions from this latter observation. The first is that the minimal requirements of Na+-driven, Disp-mediated Shh extraction and handover to Scube2 47 are insufficient to release Shh. The second is that Shh self-assembly by the law of mass action 48 is also not supported, because this process should solubilize Shh regardless of the presence or absence of serum. In contrast, we previously found that one function of Disp is to extract and transfer free plasma membrane cholesterol to a soluble sink to transport it away from the cell 4, 49. In vertebrates, HDLs represent soluble sinks for “free” peripheral cholesterol 50, and HDLs and the pharmacological cholesterol chelator methyl-β-cyclodextrin have previously been shown to increase Shh shedding 4, 23, 51. Together, these findings suggested that the permissive factor lacking from our shedding assay was most likely a soluble cholesterol acceptor, such as HDL or a related serum LPP, as reported previously 14.

Revised version: It was recently shown that the prohormone convertase furin cleaves Disp at a conserved processing site to activate it and to release Shh from the cell surface 46. Based on this mode of activation, we hypothesized that furin inhibition might specifically interfere with Disp-regulated Shh shedding. To test this hypothesis, we added 0-80 μM peptidyl chloromethylketone (CMK, a competitive inhibitor of furin) to our solubilization assays. Indeed, CMK reduced Shh shedding from the cell surface in a concentration-dependent manner (Fig. 2A-C, Fig. S3A). During these assays, we also found that repeated careful washing of cells to quantitatively remove all traces of serum severely impaired Disp- and Scube2-mediated Shh solubilization into serum-free media (Fig. 2D,E, Fig. S3B). From this latter observation, we derive two important conclusions. The first is that the minimal requirements of Na+-driven, Disp-mediated Shh extraction and hand-over to Scube2 47 are not sufficient to release Shh. The second is that Shh self-assembly by the law of mass action 48 is also not supported, because this process should solubilize Shh regardless of the presence or absence of serum. In contrast, we previously found that one function of Disp is to extract free plasma membrane cholesterol and transfer it to a soluble sink for removal from the cell 4, 49. In vertebrates, HDLs represent a soluble sink for “free” peripheral cholesterol 50, and HDLs and the pharmacological cholesterol chelator methyl-β-cyclodextrin have previously been shown to increase Shh shedding 4, 23, 51. Taken together, these results suggest that the permissive factor missing from our shedding assay is most likely a soluble cholesterol acceptor, such as HDL or a related serum LPP, as previously reported 14.

Shh shedding depends on Disp activation by furin and the presence of serum traces.

Earlier version: A) nt ctrl cells were transfected with Shh and Scube2, and Shh solubilization was monitored in the presence or absence of CMK furin inhibitor. CMK impaired proteolytic processing and release of truncated (arrowheads) soluble Shh in a concentration-dependent manner. Right: Shh release in the presence or absence of 50 μM CMK furin inhibitor. B) Quantification of CMK-inhibited Shh shedding. Ratios of solubilized versus cellular Shh in the presence of 50 μM CMK inhibitor were determined and expressed relative to Shh solubilization in the absence of inhibitor (set to 100%, black bar). Unpaired t-test, twotailed. ****: p<0.0001, n=9. See Fig. S3A for loading controls and Table S1 for additional statistical information. C) Immunoblotted cellular (c) and medium (m) fractions of Shh expressing nt ctrl and Disp-/- cells in the complete absence of serum (referred to as “serum-free” conditions throughout this work). Note that Shh solubilization under serum-free conditions is abolished.

Revised version: A) nt ctrl cells were transfected with Shh and Scube2, and Shh solubilization was monitored in the presence or absence of the furin inhibitor CMK. CMK impaired proteolytic processing and release of truncated (arrowheads) soluble Shh in a concentration-dependent manner. B) Truncated (arrowhead) Shh release in the presence or absence of 50 μM CMK furin inhibitor. C) Quantification of CMK-inhibited Shh shedding. Ratios of solubilized versus cellular Shh (set to 100%) in the presence of 50 μM CMK inhibitor were determined and expressed relative to Shh solubilization in the absence of inhibitor (DMSO). Unpaired t-test, two-tailed. **: p=0.0021, n=6. See Fig. S3A for loading controls and Table S1 for additional statistical information. D) Immunoblotted cellular (c) and medium (m) fractions of Shh expressing nt ctrl and Disp-/- cells in the complete absence of serum (referred to as “serum-free” conditions throughout this paper). Note that Shh solubilization is greatly reduced under serum-free conditions. E) Processed Shh quantifications after secretion from nt ctrl and Disp-/- cells in the presence or absence of Scube2 into serum-free medium. One-way ANOVA, Dunnett’s multiple comparisons test. **: p=0.0059, *: p=0.02, ns: p=0.54. See Table S1 for detailed statistical information.

Earlier version: Increased serum concentrations shift shedding of both Shh termini to selective N-terminal shedding

Revised version: Elevated serum concentrations shift shedding of both Shh termini to selective N-terminal shedding

Earlier version: To characterize serum-dependent shedding in more detail, we expressed dual-lipidated Shh in nt ctrl cells and Disp-/- cells and solubilized the proteins into serum-depleted Dulbecco’s Modified Eagle’s Medium (DMEM) or into DMEM supplemented with 0.05%, 5%, and 10% fetal calf serum (FCS) (Fig. 3A-D; Fig. S3C-F shows transfection and loading controls). Consistent with previous observations, we again found that Scube2 enhances Shh shedding from Disp-expressing cells into serum-depleted media (Fig. 3A, arrowhead), but does not much enhance Shh shedding from Disp-/- cells. RP-HPLC of the solubilized material confirmed that the observed increase in electrophoretic mobility was caused by the proteolytic removal of both terminal lipidated peptides (Fig. 3A’). Increased serum concentrations during Shh solubilization (0.05% FCS (Fig. 3B), 5% FCS (Fig. 3C), and 10% FCS (Fig. 3D)) did not appear to much affect Disp- and Scube2-specific Shh shedding (arrowheads). To our surprise, however, RP-HPLC of the solubilized materials shown in Fig. 3B-D revealed a gradual shift from dual Shh shedding (labeled [1] in Fig. 3A’-D’) toward the solubilization of variant proteins with their C-terminal cholesteroylated peptide still intact (labeled [2] in Fig. 3B’-D’, Fig. S4A-E). We note that low amounts of dual-lipidated Shh were also present in serum-containing media (indicated by asterisks in Fig. S4B-E). However, relative amounts of this protein fraction increased both in the absence of Scube2 (Fig. S4F; here, cells required incubation for 24 h to compensate for low Shh solubilization) and in the absence of Disp (Fig. S4G). In contrast, Scube2 expression in Disp-/- cells increased relative amounts of delipidated soluble Shh (Fig. S4H). We conclude from the Disp- and Scube2-independent solubilization of dual-lipidated Shh that it is physiologically irrelevant.

Revised version: To further characterize serum-dependent shedding, we expressed dual-lipidated Shh in nt ctrl cells and Disp-/- cells and solubilized the proteins into serum-depleted Dulbecco’s Modified Eagle’s Medium (DMEM) or into DMEM supplemented with 0.05%, 5% and 10% fetal calf serum (FCS) (Fig. 3A-D; Fig. S3C-F shows transfection and loading controls). Consistent with previous observations, we again found that Scube2 enhanced Shh shedding from Disp-expressing cells into serum-depleted media (Fig. 3A, arrowhead), but did not significantly enhance Shh shedding from Disp-/- cells. RP-HPLC of the solubilized material confirmed that the observed increase in electrophoretic mobility was caused by the proteolytic removal of both terminal lipidated peptides (Fig. 3A’’). Increased serum concentrations during Shh solubilization (0.05% FCS (Fig. 3B, quantified in B’), 5% FCS (Fig. 3C, quantified in C’), and 10% FCS (Fig. 3D, quantified in D’)) did not appear to much affect Disp- and Scube2-specific Shh shedding (arrowheads, Fig. 3E in contrast shows unregulated release of unprocessed Shh (the “top” band) in the presence of 10% FCS). To our surprise, however, RP-HPLC of the solubilized materials shown in Fig. 3B-D revealed a gradual shift from dual Shh shedding (labeled [1] in Fig. 3A’’-D’’) toward the solubilization of variant proteins with their C-terminal cholesteroylated peptide still intact (labeled [2] in Fig. 3B’’-D’’, Fig. S4A-E). We also observed this hydrophobicity shift (together with an electrophoretic mobility shift, see Fig. S1D) when analyzing endogenous Shh from the pancreatic cancer cell line Panc1 (Fig. S4F-H). We note that small amounts of dual-lipidated Shh were also present in serum-containing media (indicated by asterisks in Fig. S4B-E). However, the relative amounts of this protein fraction increased both in the absence of Scube2 (Fig. S4I; here, cells had to be incubated for 24 h to compensate for the low Shh solubilization) and in the absence of Disp (Fig. S4J). In contrast, Scube2 expression in Disp-/- cells increased the relative amount of delipidated soluble Shh (Fig. S4K). From this, we conclude that the Disp- and Scube2-independent solubilization of dual-lipidated Shh is physiologically irrelevant, consistent with the findings shown in Fig. F2F-H..

Earlier version: Dual-lipidated cell-surface Shh converts into delipidated soluble forms under low-serum and high-serum conditions.

Revised version: Dually lipidated cell surface Shh converts to delipidated soluble forms under low-serum and high-serum conditions.

Earlier version: Media were changed for serum-free DMEM 36 h post-transfection (cells were not washed) or DMEM containing the indicated amounts of serum, and proteins were solubilized for 6 h (serum-depleted) or for 24 h (with 0.05%, 5%, and 10% serum). A) Under serum-depleted conditions, Disp and Scube2 increase dual-lipidated Shh shedding into truncated soluble forms (arrowhead). m: media, c: cell lysate. A’) RP-HPLC confirmed loss of both terminal lipidated Shh peptides during Disp- and Scube2-regulated shedding. B-D) The appearance of truncated Shh in serum-containing media remained dependent on Disp and – to a lesser degree – on Scube2 (arrowheads). B’-D’) RP-HPLC revealed that increased serum amounts shift dual Shh shedding (low serum amounts, [1]) toward N-terminally restricted shedding and release of a cholesteroylated Shh form [2]. Low levels of dual-lipidated Shh are also detected [3]. See Fig. S4 for experimental details.

Revised version: Media were changed to serum-free DMEM (cells were not washed) or DMEM containing the indicated amounts of serum 36 h post-transfection, and proteins were solubilized for 6 h (serum-depleted) or for 24 h (with 0.05%, 5%, and 10% serum). A) Under serum-depleted conditions, Disp and Scube2 increase the conversion of dual-lipidated Shh into truncated soluble forms (arrowhead). m: media, c: cell lysate. A’’) RP-HPLC confirmed the loss of both terminal lipidated Shh peptides during Disp- and Scube2-regulated shedding, as shown previously (Fig. 1A’’). B-D) The appearance of truncated Shh in serum-containing media remained dependent on Disp and, to a lesser extent, Scube2 (arrowheads). B’-D’) Quantifications of B-D. One-way ANOVA, Dunnett’s multiple comparisons test. ****: p<0.0001, ***: p<0.001, **: p<0.01, *: p<0.05. See Table S1 for detailed statistical information. B’’-D’’) RP-HPLC revealed that increased serum levels shift dual Shh shedding (low serum levels, [1]) towards N-terminally restricted shedding and release of a cholesteroylated Shh form [2]. Low levels of dually lipidated Shh are also detected [3]. E) Quantification of unprocessed Shh release in the presence of 10% FCS (top band in D). One-way ANOVA, Dunnett’s multiple comparisons test. ns: p>0.05. See Fig. S4 for additional information.

Earlier version: So far, our data showed that Scube2 increases N-terminal and C-terminal Shh shedding from Disp-expressing cells and that serum enhances this process, likely by providing a sink for membrane cholesterol transferred by Disp 4 (Fig. S4I). We also showed that high serum amounts promote a second Disp solubilization mode in which Shh shedding is restricted to the palmitoylated N-peptide while the C-terminus remains intact. This latter finding raised the interesting possibility that serum factors may accept or protect the cholesteroylated Shh C-terminus. Indeed, it is known that a filtrate of blood serum through the capillary walls, called interstitial fluid, represents the microenvironment of tissues and cells in vivo, as well as of those that express and solubilize Shh during development 14. It is also known that interstitial fluid is rich in proteins and LPP particles of small mass from the serum 52. This suggests that Shh expression in the presence of serum may resemble the conditions in Shh-expressing tissues in vivo, and may therefore be closer to physiological relevance.

Revised version: So far, our data have shown that Scube2 increases N- and C-terminal Shh shedding from Disp-expressing cells and that serum enhances this process, probably by providing a sink for membrane cholesterol transferred by Disp 4 (Fig. S4L). We also showed that high serum levels promote a second Disp solubilization mode in which Shh shedding is restricted to the palmitoylated N-peptide, leaving the C-terminus intact. This latter finding raised the interesting possibility that serum factors may accept or protect the cholesteroylated Shh C-terminus. Indeed, it is known that a filtrate of blood serum through the capillary walls, called interstitial fluid, represents the microenvironment of tissues and cells in vivo, as well as of tissues and cells that express and solubilize Shh during development 14. It is also known that the interstitial fluid is rich in LPPs of small mass from the serum 52. This suggests that Shh expression in the presence of serum may resemble the conditions in Shh-expressing tissues in vivo and thus be closer to physiological relevance.

Earlier version: N-palmitate is not required for Shh-induced in vitro differentiation of C3H10T1/2 cells

Revised version: Shh-induced in vitro differentiation of C3H10T1/2 cells does not require N-palmitate

Earlier version: Is this novel N-terminally processed Shh variant fully functional? It is well-established that dual lipidation during biosynthesis is absolutely required for unperturbed Hh/Shh biofunction in vivo 39, 53, 54, 55, 56. These studies showed that dual-lipidated Hh/Shh expression in vitro generates soluble variants that are 10-30 times more bioactive than artificial proteins that do not undergo Hhat-catalyzed N-palmitoylation, yet undergo unperturbed C-terminal cholesteroylation and secretion to the cell surface (Fig. 1B) 42, 56. According to these published observations, the N-terminally processed Shh variant described here should not be very active because it also lacks the palmitate. To test this hypothesis, we used the Ptch1-expressing, multipotent fibroblastic C3H10T1/2 cell line as a reporter 57. We first verified multipotency of our C3H10T1/2 reporter cell line to differentiate into osteoblasts 57, chondrocytes 58, or adipocytes 59 (Fig. S5A). To this end, C3H10T1/2 cells were cultured for various time periods in the presence of adipogenic, chondrogenic, and osteogenic supplements, and their responsiveness was confirmed on the basis of phenotype and the expression of cell surface markers. We then incubated C3H10T1/2 cells with Shh expressed in the presence of 10% serum and physiological release regulators Scube2 30, 31 and Disp 5, 6 (Fig. 4A). Shh shares 91% sequence identity and both lipids with Indian hedgehog, an established osteogenic factor 60, and both Indian hedgehog and Shh stimulate C3H10T1/2 osteogenic differentiation. We confirmed that Shh induced alkaline phosphatase (Alp) expression and C3H10T1/2 differentiation into osteoblasts in a concentration-dependent manner 57 (Fig. 4A). Notably, we observed that C25AShh expressed under the same conditions was equally active or even more active than Shh (Fig. 4A). Quantitative reverse transcription-polymerase chain reaction analysis (qPCR) of Ptch1 and Gli1 expression in Shh-treated C3H10T1/2 cells (Fig. 4B, Fig. S5B) or NIH3T3 cells (Fig. S5B) confirmed this finding: After protein normalization, qPCR confirmed similar increases in mRNA expression of Ptch1 and Gli1 under direct Shh or C25AShh control. Ptch1 is known to be upregulated by Shh 61, and Gli1 is a zinc finger transcription factor that operates downstream of Ptch1 and is also transcribed in a Hh-dependent manner 62. Key regulators for adipogenesis, osteogenesis, chondrogenesis, and proliferation in Shh and C25SShh-induced C3H10T1/2 cells were also similar (Fig. S5C,D). In contrast, C25SShhN negative control biofunction was always strongly impaired. Finally, we incubated C3H10T1/2 cells with high-activity Shh R&D 8908-SH, a commercially available, dual-lipidated Shh variant obtained by detergent extraction from transfected cells (Fig. S5E). We observed that R&D 8908-SH induced Alp expression in differentiating C3H10T1/2 cells in a concentration-dependent manner, as expected 57. We also found that the activities of solubilized Shh and C25AShh were increased over similar amounts of R&D 8908-SH. qPCR of Ptch1 and Gli1 expression confirmed high activities of R&D 8908-SH, Shh and C25AShh (Fig. S5F). These results demonstrate that N-palmitate is not essential for Shh signaling strength at the level of Ptch1 if Shh is released in the presence of serum.

Revised version: Is this novel N-terminally processed Shh variant fully functional? It is well established that dual lipidation during biosynthesis is absolutely necessary for unimpaired Hh/Shh biofunction in vivo 39, 53, 54, 55, 56. These studies have shown that dual-lipidated Hh/Shh expression in vitro generates soluble variants that are 10-30 times more bioactive than engineered proteins that do not undergo Hhat-catalyzed N-palmitoylation but undergo unperturbed C-terminal cholesteroylation and secretion to the cell surface (Fig. 1B) 42, 56. According to these published observations, the N-terminally processed Shh variant described here should not be very active because it also lacks the palmitate. To test this hypothesis, we used the Ptch1-expressing multipotent fibroblastic cell line C3H10T1/2 as a reporter 57. We first verified the multipotency of our C3H10T1/2 cells to differentiate into osteoblasts 57, chondrocytes 58, or adipocytes 59 (Fig. S5A). To this end, C3H10T1/2 cells were cultured in the presence of adipogenic, chondrogenic, and osteogenic supplements for different periods of time, and their responsiveness was confirmed based on phenotype and the expression of cell surface markers. We then incubated C3H10T1/2 cells with Shh expressed in the presence of 10% serum and the physiological release regulators Scube2 30, 31 and Disp 5, 6 (Fig. 4A). Shh shares 91% sequence identity and both lipids with Indian hedgehog, an established osteogenic factor 60, and both Indian hedgehog and Shh stimulate C3H10T1/2 osteogenic differentiation. We confirmed that Shh induced alkaline phosphatase (Alp) expression and C3H10T1/2 differentiation into osteoblasts in a concentration-dependent manner 57 (Fig. 4A). Notably, we observed that C25AShh expressed under the same conditions was as active or even more active than Shh (Fig. 4A). Quantitative reverse transcription-polymerase chain reaction (qPCR) analysis of Ptch1 and Gli1 expression in Shh-treated C3H10T1/2 cells (Fig. 4B, Fig. S5B) or Shh-treated NIH3T3 cells (Fig. S5B) confirmed this finding: After protein normalization, qPCR confirmed similar increases in Ptch1 and Gli1 mRNA expression under direct Shh or C25AShh control. Ptch1 is known to be upregulated by Shh 61, and Gli1 is a zinc finger transcription factor that acts downstream of Ptch1 and is also transcribed in an Hh-dependent manner 62. Key regulators of adipogenesis, osteogenesis, chondrogenesis, and proliferation in Shh and C25SShh-induced C3H10T1/2 cells were also similar (Fig. S5C,D). In contrast, the biofunction of the C25SShhN negative control was always severely impaired. Finally, we incubated C3H10T1/2 cells with high activity Shh R&D 8908-SH, a commercially available dual lipidated Shh variant obtained by detergent extraction from transfected cells (Fig. S5E,F). We observed that R&D 8908-SH induced Alp expression in differentiating C3H10T1/2 cells in a concentration-dependent manner, as expected 57. We also found that the activities of solubilized Shh and C25AShh were increased over similar amounts of R&D 8908-SH. qPCR of Ptch1 and Gli1 expression confirmed high activities of R&D 8908-SH, Shh and C25AShh (Fig. S5G). These results demonstrate that N-palmitate is not essential for the strength of Shh signaling at the level of Ptch1 when C25AShh is released in the presence of serum.

Added in the revised version: Activities and SEC of dual-lipidated Shh/Hh and depalmitoylated Shh variants solubilized into serum-containing media.

Earlier version: Activities and SEC of dual-lipidated Shh/Hh and depalmitoylated Shh variants solubilized into serum-containing media. A) Shh, C25AShh (this artificial variant has the cysteine palmitate acceptor changed for a non-accepting alanine), and non-lipidated C25AShhN were expressed into media containing 10% FCS; their protein amounts determined by immunoblotting and normalized; and conditioned media added to C3H10T1/2 reporter cells to induce their Hh-dependent differentiation into alkaline phosphatase (Alp)-producing osteoblasts. Mock-treated C3H10T1/2 cells served as non-differentiating controls. At lower concentrations (1x), Shh and C25AShh induced C3H10T1/2 differentiation in a similar manner, as determined by Alp activity measured at 405 nm. At higher concentrations (2x), the bioactivity of C25AShh was increased over that of Shh. C25AShhN was inactive. One-way ANOVA, Sidaks multiple comparisons test. ****: p<0.0001, n.s.=0.99, n=3-9. See Table S1 for additional statistical information and Fig. S5E,F, which demonstrates protein activities similar to, or exceeding, that of a dual-lipidated Shh control. B) Similar transcription of Hh target genes Ptch1 and Gli1 by Shh and C25AShh in two independent experiments. C3H10T1/2 reporter cells were stimulated with similar amounts of Shh, C25AShh, and C25AShhN at high (2x) and lower (1x) concentrations. C) SEC reveals substantial amounts of increased molecular weight Shh in media that contain 10% serum (black line). The increased molecular weight Shh eluted together with ApoA1 (orange line). In contrast, Scube2 was largely monomeric in solution (blue line). Levels of Shh-induced Alp activity in C3H10T1/2 cells were measured as absorption at 405 nm, demonstrating the strongest C3H10T1/2 differentiation by eluted fractions containing large Shh assemblies. D) SEC of Drosophila Hh (black line) and of a variant lacking its HS binding site (HhβHS, black dotted line). Both proteins were expressed under actin-Gal4/UAS-control from S2 cells and solubilized into media containing 10% FCS.

Revised version: A) Shh, C25AShh (this artificial variant has the cysteine palmitate acceptor changed to a non-accepting alanine), and non-lipidated C25AShhN were expressed in media containing 10% FCS; their protein levels were determined by immunoblotting and normalized; and the conditioned media were added to C3H10T1/2 reporter cells to induce their Hh-dependent differentiation into alkaline phosphatase (Alp)-producing osteoblasts. Mock-treated C3H10T1/2 cells served as non-differentiating controls. At lower concentrations (1x), Shh and C25AShh induced C3H10T1/2 differentiation in a similar manner, as determined by Alp activity measured at 405 nm. At higher concentrations (2x), the bioactivity of C25AShh was increased over that of Shh. C25AShhN was inactive. One-way ANOVA, Sidaks multiple comparisons test. ****: p<0.0001, ns=0.99, n=3-9. See Table S1 for additional statistical information and Fig. S5E,F, which shows protein activities similar to or exceeding those of a dual-lipidated Shh control. B) Similar transcription of the Hh target genes Ptch1 and Gli1 by Shh and C25AShh in three independent experiments. C3H10T1/2 reporter cells were stimulated with similar amounts of Shh, C25AShh, and C25AShhN at high (2x) and low (1x) concentrations. C) SEC shows significant amounts of Shh of increased molecular weight in media containing 10% serum (black line). The increased molecular weight Shh eluted together with ApoA1 (orange line). In contrast, Scube2 was largely monomeric in solution (blue line). The level of Shh-induced Alp activity in C3H10T1/2 cells was measured as absorbance at 405 nm, showing the strongest C3H10T1/2 differentiation by eluted fractions containing large Shh aggregates. D) SEC of Drosophila Hh (black line) and of a variant lacking its HS binding site (HhΔHS, black dotted line). Both proteins were expressed from S2 cells under actin-Gal4/UAS-control and solubilized into media containing 10% FCS.

Earlier version: Soluble Shh/Hh and ApoA1-containing LPPs have similar sizes

Revised version: Soluble Shh/Hh and ApoA1-containing LPPs have a similar size

Earlier version: A decade ago, it was already known that flies and mammals release sterol-modified Hh/Shh in LPP-associated bioactive form, yet also in desteroylated, non-associated form 14. These findings suggested that serum LPPs in our assays not only may have promoted dual Shh shedding 4 (Fig. 1), but may also have promoted monolipidated Shh assembly into soluble complexes 14. We used size exclusion chromatography (SEC) to test this possibility. SEC detects soluble Shh monomers (20 kDa) and covers the entire molecular weight (MW) range up to a cutoff of about 1 mDa. The MW span of serum HDL ranges from 1.75×105 Da to 3.6×105 Da, corresponding to particle sizes of 5-11 nm. This small size range makes HDLs abundant constituents of interstitial fluid and good candidates for Shh solubilization from cells and tissues. Larger complexes in mammalian serum include low-density lipoproteins (LDLs, MW 2.75×106 Da) and very low-density lipoproteins (VLDLs, MW 10-80×106 Da) that would both elute in the void column volume. As shown in Fig. 4C, Shh in DMEM + 10% FCS elutes in fractions that cover the entire MW range (black line). One prominent Shh peak is detected around 20 kDa – most likely representing the fully delipidated dually shed Shh fraction (Fig. 3, labeled [1]) – and a second peak is detected between 300 kDa and 600 kDa. Reprobing the same (stripped) blot with antibodies directed against ApoA1 (the major protein component of HDL) revealed an elution profile that overlapped with that of Shh (Fig. 4C, orange line). The leftward shift of Shh elution relative to HDL elution may be explained by the increased size of a LPP subfraction after its Disp-mediated loading with Shh. In contrast, Scube2 (Fig. 4C, blue line) co-eluted only with smaller Shh multimers 48. In order to determine which size fraction contained biologically active Shh, we analyzed C3H10T1/2 osteoblast precursor cell differentiation by using aliquots of eluted fractions from the same SEC run. We found that Shh induced C3H10T1/2 differentiation into osteoblasts in a size-dependent activity distribution: Large Shh assemblies were strongly bioactive, and smaller assemblies were also bioactive, but monomeric Shh was only moderately active (Fig. 4C, green dashed line). We also found that Drosophila Hh expressed in S2 cells and solubilized into serum-containing medium showed a similar size distribution (Fig. 4D, black line) and that Hh assemblies also co-eluted with serum ApoA1 (orange line). This suggests that the mechanism of serum-enhanced Hh assembly during solubilization is conserved. The alternative possibility of Hh assembly as a consequence of heparan sulfate (HS) proteoglycan interactions 63 was dismissed because site-directed mutagenesis of the HS binding Cardin-Weintraub motif did not change HhΔHS size distribution (Fig. 4D, black dotted line). These findings confirm a previous study showing that Drosophila Hh copurifies with LPP and colocalizes with LPP in the developing wing epithelium in vivo 15. The same study showed that reduced LPP levels in larvae impair Hh release and lead to Hh accumulation at the site of production, similar to the in vitro findings shown in Fig. 2C.

Revised version: A decade ago, it was already known that flies and mammals release sterol-modified Hh/Shh in the LPP-associated bioactive form as well as in the desteroylated, unassociated form 14. These results suggest that serum LPPs in our assays may not only have promoted dual Shh shedding 4 (Fig. 1), but may also have promoted assembly of monolipidated Shh into soluble LPP-associated complexes 14. We used size exclusion chromatography (SEC) to test this possibility. SEC detects soluble Shh monomers (20 kDa) and covers the entire molecular weight (MW) range up to a cutoff of about 1 mDa. The MW range of serum HDL is from 1.75×105 Da to 3.6×105 Da, corresponding to particle sizes of 5-11 nm. This small size range makes HDLs abundant components of interstitial fluid and good candidates for Shh solubilization from cells and tissues. Larger complexes in mammalian serum include low-density lipoproteins (LDLs, MW 2.75×106 Da) and very low-density lipoproteins (VLDLs, MW 10-80×106 Da), both of which would elute in the void column volume. As shown in Fig. 4C, Shh in DMEM + 10% FCS elutes in fractions covering the entire MW range (black line). A prominent Shh peak is detected at around 20 kDa – most likely representing the dually cleaved, fully delipidated Shh fraction (as shown in Fig. 3A’’-D’’, labeled [1]) – and a second peak is detected between 300 kDa and 600 kDa. Reprobing the same (stripped) blot with antibodies directed against ApoA1 (the major protein component of HDL) revealed an elution profile that overlapped with that of Shh (Fig. 4C, orange line). The leftward shift of Shh elution relative to HDL elution may be explained by the increased size of an LPP subfraction after its Disp-mediated loading with Shh. In contrast, Scube2 (Fig. 4C, blue line) co-eluted only with smaller Shh multimers 48. To determine which size fraction contained biologically active Shh, we analyzed the differentiation of C3H10T1/2 osteoblast progenitor cells by using aliquots of eluted fractions from the same SEC run. We found that Shh induced C3H10T1/2 osteoblast differentiation with a size-dependent activity distribution: Large Shh assemblies were highly bioactive, and smaller assemblies were also bioactive, but monomeric Shh was only moderately active (Fig. 4C, green dashed line). We also found that Drosophila Hh expressed in S2 cells and solubilized into serum-containing medium showed a similar size distribution (Fig. 4D, black line) and that Hh assemblies also co-eluted with serum ApoA1 (orange line). This suggests that the mechanism of serum-enhanced Hh assembly during solubilization is conserved. The alternative possibility of Hh assembly as a consequence of heparan sulfate (HS)-proteoglycan interactions 63 was rejected because site-directed mutagenesis of the HS-binding Cardin-Weintraub motif did not alter the size distribution of HhΔHS (Fig. 4D, black dotted line). These results confirm a previous study showing that Drosophila Hh copurifies with LPP and co-localizes with LPP in the developing wing epithelium in vivo 15. The same study showed that reduced LPP levels in larvae impair Hh release and lead to Hh accumulation at the site of production, similar to the in vitro results shown in Fig. 2D,E.

Disp function requires HDL

Earlier version: To test whether the soluble LPP that increases Disp-mediated solubilization of N-processed, cholesteroylated Shh is HDL, we tested this possibility directly. To this end, we expressed Shh in nt ctrl cells and Disp-/- cells in the presence or absence of Scube2, washed the cells 36 h post-transfection to inhibit Shh release, and added serum-free DMEM or DMEM supplemented with 40 μg/mL purified human HDL 4. We first confirmed that lack of serum factors strongly reduced Shh solubilization (Fig. 5A, labeled [1], Fig. S6A). RP-HPLC of the small amounts of released Shh confirmed removal of both lipidated Shh peptide anchors during solubilization (Fig. 5A’, Fig. S6A’), and SEC confirmed that solubilized proteins were monomeric (Fig. 5A’’). In contrast, Shh solubilization from nt ctrl cells into DMEM supplemented with 40 μg/mL HDL was strongly increased (Fig. 5B,C, Fig. S6B). As previously observed for Shh expression in the presence of serum, RP-HPLC of HDL-solubilized Shh (Fig. 5B, labeled [2]) confirmed that its C-terminal cholesteroylated peptides were still intact (Fig. 5B’, Fig. S6B’) and proteins were strongly bioactive (Fig. 5B’’, green dashed line). Solubilized dual-lipidated Shh was also found, but its solubilization was again independent of Scube2 and Disp function (Fig. S6B’’,B’’’), indicating unspecific release. From the observed size shift and overlap with ApoA1 and ApoE4 elution profiles in SEC (note that the observed Shh MW range almost matches that of ApoA1), we suggest that monolipidated truncated Shh associates with HDL in a Disp-regulated manner. This possibility is supported by the previous observation that Hh levels in the hemolymph of Disp-deficient fly larvae were strongly decreased 14. We also tested whether HDL constitutes the only possible Shh acceptor, or whether LDL can also carry monolipidated morphogens. As shown in Fig. S6C,D, Shh transferred to HDL-supplemented serum-free DMEM, but not to LDL-supplemented serum-free DMEM. Shh released by HDL is N-terminally processed (Fig. S6G) and physically interacts with the LPP (Fig. S6E,F).

Revised version: To test whether the soluble LPP that enhances Disp-mediated solubilization of N-processed, cholesteroylated Shh is HDL, we tested this possibility directly. To this end, we expressed Shh in nt ctrl cells and Disp-/- cells in the presence or absence of Scube2, washed the cells 36 h after transfection to inhibit Shh release, and added serum-free DMEM or DMEM supplemented with 40 μg/mL purified human HDL 4. We first confirmed that the absence of serum factors greatly reduced Shh solubilization (Fig. 5A, labeled [1], Fig. S6A). RP-HPLC of the small amounts of released Shh confirmed the removal of both lipidated Shh peptide anchors during solubilization (Fig. 5A’, Fig. S6A’), and SEC confirmed that the solubilized proteins were monomeric and not Scube2-associated (Fig. 5A’’). In contrast, Shh solubilization from nt ctrl cells into DMEM supplemented with 40 μg/mL HDL was greatly increased (Fig. 5B,C, Fig. S6B). As previously observed for Shh expression in the presence of serum, RP-HPLC of HDL-solubilized Shh (Fig. 5B, labeled [2]) confirmed that its C-terminal cholesteroylated peptides were still intact (Fig. 5B’, Fig. S6B’) and that the proteins were highly bioactive (Fig. 5B’’, green dashed line). Solubilized dual-lipidated Shh – represented by the “upper” bands – was also found, but its solubilization was again independent of Scube2 and Disp function (Fig. S6B’’,B’’’), indicating a non-specific release. From the observed size shift and overlap with ApoA1 and ApoE4 elution profiles in SEC (Fig. 5B’’, note that the observed Shh MW range almost matches that of ApoA1), we suggest that monolipidated truncated Shh associates with HDL in a Disp-regulated manner. This possibility is supported by the previous observation that Hh levels were greatly decreased in the hemolymph of Disp-deficient fly larvae 14. We also tested whether HDL is the only possible Shh acceptor, or whether LDL can also carry monolipidated morphogens. As shown in Fig. S6C,D, Shh was transferred to HDL-supplemented serum-free DMEM, but not to LDL-supplemented serum-free DMEM. Shh released by HDL is N-terminally processed (Fig. S6G) and physically interacts with the LPP (Fig. S6E,F).

Earlier version: HDL increases N-processed Shh solubilization by Disp.

Revised version: HDL enhancess N-processed Shh solubilization by Disp.

Earlier version: Media were changed for serum-free DMEM 36 h post-transfection (and cells washed three times) or serum-free DMEM supplemented with 40 μg/mL HDL before proteins were solubilized for 6 h. A) Immunoblotted cellular (c) and medium (m) fractions of Shh expressing nt ctrl and Disp-/- cells in the complete absence of serum. Note that Shh solubilization is strongly reduced [1]. A’) RP-HPLC of the material labeled [1] revealed complete delipidation during release. A’’) SEC of the same delipidated material demonstrates that it is readily soluble and not associated with Scube2. B) Immunoblotted cellular (c) and medium (m) fractions of Shh expressing nt ctrl and Disp-/- cells in the presence of 40 μg/mL HDL. Shh shedding and solubilization are strongly increased by Disp [2] but not by Scube2. B’) RP-HPLC of the material labeled [2] revealed that HDL shifts Shh shedding from dual processing (A’, [1]) toward the release of cholesteroylated Shh. B’’) SEC of the same material [2] (black line) shows a molecular weight increase that matches the molecular weight range of HDL, as indicated by marker apolipoproteins ApoA1 (orange line) and mobile ApoE4. The former lends structural stability to the particle and stimulates cholesterol efflux to HDL; the latter facilitates cholesterol storage and core expansion and is therefore a marker of larger mature HDL particles (brown dashed line). Again, the soluble Shh elution profile does not overlap with that of Scube2 (blue line). C) Quantification of HDL-induced Shh solubilization from nt ctrl cells and Disp-/- cells. Unpaired t-test, two-tailed. ****: p<0.0001, n=6. See Table S1 for additional statistical information.

Revised version: Media were changed to serum-free DMEM (and cells washed three times) or serum-free DMEM supplemented with 40 µg/mL HDL 36 h post-transfection before protein solubilization for 6 h. A) Immunoblotted cellular (c) and medium (m) fractions of Shh expressing nt ctrl and Disp-/- cells in the complete absence of serum. Note that Shh solubilization is greatly reduced [1]. A’) RP-HPLC of the material labeled [1] in A showed complete delipidation during release. A’’) SEC of the same delipidated material shows that it is readily soluble and not associated with Scube2. B) Immunoblotted cellular (c) and medium (m) fractions of Shh expressing nt ctrl and Disp-/- cells in the presence of 40 μg/mL HDL. Shh shedding and solubilization are strongly increased by Disp [2] but not by Scube2. B’) RP-HPLC of the material labeled [2] in B showed that HDL shifts Shh shedding from dual processing (A’, [1]) to release of cholesteroylated Shh. B’’) SEC of the same material [2] (black line) shows an increase in molecular weight corresponding to the molecular weight range of HDL, as indicated by the marker apolipoproteins ApoA1 (orange line) and mobile ApoE4. The former provides structural stability to the particle and stimulates cholesterol efflux to HDL; the latter facilitates cholesterol storage and core expansion and is therefore a marker of larger mature HDL particles (brown dashed line). Again, the soluble Shh elution profile does not overlap with that of Scube2 (blue line). C) Quantification of HDL-induced Shh solubilization from nt ctrl cells and Disp-/- cells. One-way ANOVA, Dunnett’s multiple comparisons test. ***: p=0.0008, **: p=0.0023, ns: p=0.77, n=7. Additional statistical information is provided in Table S1.

Earlier version: Is N-terminally processed, HDL-associated Shh functional? To answer this question, we again used the multipotent fibroblastic C3H10T1/2 cell line as a reporter. Shh was co-expressed with Scube2 in the presence of 80 μg/mL purified human HDL and the conditioned media added to C3H10T1/2 cells. As expected, Shh/HDL complexes induced Alp expression and C3H10T1/2 differentiation into osteoblasts in a concentration-dependent manner 57, and C25AShh expressed under the same conditions was equally active (Fig. 6A). qPCR of Ptch1 and Gli1 expression in Shh/HDL-stimulated or C25AShh/HDL-stimulated C3H10T1/2 cells confirmed this finding: Both protein/HDL complexes induced similar increases in mRNA expression of Ptch1 and Gli1 (Fig. 6B). This suggests full bioactivity of the HDL-associated, N-terminally truncated Shh variants that we describe in this study.

Revised version: Is N-terminally processed, HDL-associated Shh functional? To answer this question, we again used the multipotent fibroblastic cell line C3H10T1/2 as a reporter. Shh was co-expressed with Scube2 in the presence of 80 µg/mL purified human HDL and the conditioned media was added to C3H10T1/2 cells. As expected, Shh/HDL complexes induced Alp expression and C3H10T1/2 differentiation into osteoblasts in a concentration-dependent manner 57 and C25AShh expressed under the same conditions was equally active (Fig. 6A). qPCR of Ptch1 and Gli1 expression in Shh/HDL-stimulated or C25AShh/HDL-stimulated C3H10T1/2 cells and NIH3T3 cells confirmed this finding: Both protein/HDL complexes induced similar increases in Ptch1 and Gli1 mRNA expression (Fig. 6B). This suggests an unchanged bioactivity of the HDL-associated, N-terminally truncated Shh variants described in this study.

Added in the revised version: Activities of HDL-associated Shh and non-palmitoylated variants.

Earlier version: Activities of HDL-associated Shh and non-palmitoylated variants. A) Shh, non-palmitoylated C25AShh, and non-lipidated C25AShhN were released into media containing 80 μg/mL HDL; their protein amounts determined by immunoblotting and normalized; and conditioned media added to C3H10T1/2 reporter cells to induce their differentiation into Alp-producing osteoblasts. Mock-treated C3H10T1/2 cells served as non-differentiating controls. At low (1x) and high (2x) concentrations, Shh and C25AShh induced C3H10T1/2 differentiation in a similar manner, as determined by Alp activity measured at 405 nm. Again, C25AShhN was completely inactive, in contrast to bioactive HDL-associated non-palmitoylated C25SShh. One-way ANOVA, Sidaks multiple comparisons test. ****: p<0.0001, ***: p<0.001, n.s. >0.1, n=4. See Table S1 for additional statistical information. B) Similar transcription of Hh target genes Ptch1 and Gli1 by HDL-associated Shh and C25AShh in two independent experiments. C3H10T1/2 reporter cells were stimulated with similar amounts of Shh, C25AShh, and C25AShhN at high (2x) and half (1x) concentrations.

Revised version: A) Shh, non-palmitoylated C25AShh, and non-lipidated C25AShhN were released into media containing 80 µg/mL HDL, their protein levels were determined by immunoblotting and normalized, and conditioned media were added to C3H10T1/2 reporter cells to induce their differentiation into Alp-producing osteoblasts. Mock-treated C3H10T1/2 cells served as non-differentiating controls. At low (1x) and high (2x) concentrations, Shh and C25AShh induced C3H10T1/2 differentiation in a similar manner, as determined by Alp activity measured at 405 nm. Again, C25AShhN was completely inactive, in contrast to the bioactive HDL-associated non-palmitoylated C25SShh. One-way ANOVA, Sidaks multiple comparisons test. ****: p<0.0001, ***: p<0.001, ns >0.1, n=4. Additional statistical information is provided in Table S1. B) Similar transcription of Hh target genes Ptch1 and Gli1 by HDL-associated Shh and C25AShh in C3H10T1/2 cells three independent experiments. Experiment 4 confirms similar Shh and C25AShh activities in NIH3T3 cells. Reporter cells were stimulated with similar amounts of Shh, C25AShh, and C25AShhN at high (2x) and low (1x) concentrations as determined by immunoblotting.

Earlier version: The C-terminal cholesteroylated Shh peptide is required and sufficient for Disp-mediated protein export and HDL association

Revised version: The C-terminal cholesterylated Shh peptide is necessary and sufficient for the Disp-mediated export of the protein and the association with HDL

Earlier version: We asked next whether both terminal lipids are required for Disp-mediated Shh export to HDL, or whether the C-cholesteroylated peptide is sufficient for the transfer. The latter possibility was suggested by N-palmitoylated peptide loss during solubilization from Dispexpressing cells in the presence of serum (Figs. 1-3) and by Disp-independent ShhN solubilization into serum-depleted medium (Fig. 1C) 64. We expressed palmitoylated/noncholesteroylated ShhN and non-palmitoylated/cholesteroylated C25SShh and confirmed that palmitoylated ShhN solubilization from the plasma membrane was independent of Disp and HDL (Fig. 7A, [1] denotes material solubilized into serum-free DMEM, [2] denotes material solubilized by HDL; see also Fig. S7A,B). SEC detected only monomeric ShhN, regardless of the absence or presence of HDL (Fig. 7A’,A’’). These findings suggest that the N-terminal Shh palmitate is not a substrate for Disp-mediated transfer to HDL. In contrast, only small amounts of monomeric C25SShh were solubilized into serum-free media (Fig. 7B,B’, labeled [3]), and HDL strongly increased soluble C25SShh and its assembly into sizes similar to those of HDL (Fig. 7B’’, labeled [4]; see also Fig. S7C,D). Supporting the importance of Shh cholesteroylation for its transfer to HDL, soluble C25SShh C-termini remained lipidated (Fig. S7E). Together, our findings suggest that the C-terminal cholesterol moiety is required for Disp-mediated Shh transfer and HDL association. However, we also noted that a fraction of monolipidated C25SShh solubilized in a Disp-independent manner (Fig. 7B, asterisks; Fig. S7F, Fig. 7C; the green dotted line indicates the distribution of C25SShh sizes if released from Disp-/- cells). We explain this finding by unspecific C25SShh release in the presence of HDL, consistent with increased desorption of large proteins – such as Shh – if attached only to a single membrane anchor 3. Once desorbed, lipidated proteins can either reintegrate into the plasma membrane or associate with lipophilic soluble competitors, as shown in Fig. S7G. This demonstrates that another important role of fully conserved dual N- and C-terminal Hh lipidation is to reduce non-enzymatic morphogen desorption from the plasma membrane.

Revised version: We next asked whether both terminal lipids are required for Disp-mediated Shh export to HDL, or whether the C-cholesteroylated peptide is sufficient for the transfer. The latter possibility was suggested by the loss of the N-palmitoylated peptide during solubilization from Disp-expressing cells in the presence of serum (Figs. 1-3) and by Disp-independent ShhN solubilization into serum-depleted medium (Fig. 1C) 64. We expressed palmitoylated/non-cholesterylated ShhN and non-palmitoylated/cholesterylated C25SShh and confirmed that palmitoylated ShhN solubilization from the plasma membrane was independent of Disp and HDL (Fig. 7A, [1] denotes material solubilized into serum-free DMEM, [2] denotes material solubilized by HDL; see also Fig. S7A,B). SEC detected only monomeric ShhN, regardless of the presence or absence of HDL (Fig. 7A’,A’’). These results suggest that the N-terminal palmitate of Shh is not a substrate for Disp-mediated transfer to HDL. In contrast, only small amounts of monomeric C25SShh were solubilized into serum-free media (Fig. 7B,B’, labeled [3]), and HDL strongly increased soluble C25SShh and its assembly to sizes similar to those of HDL (Fig. 7B’’, labeled [4]; see also Fig. S7C,D). In support of the importance of Shh cholesterylation for its transfer to HDL, the C-termini of soluble C25SShh remained lipidated (Fig. S7E). Taken together, our results suggest that the C-terminal cholesterol moiety is required for Disp-mediated Shh transfer and HDL association. However, we also found that a fraction of monolipidated C25SShh was solubilized in a Disp-independent manner (Fig. 7B, asterisks; Fig. S7F, Fig. 7C; the green dotted line indicates the distribution of C25SShh sizes when released from Disp-/- cells). We explain this finding by the unspecific release of C25SShh in the presence of HDL, consistent with increased desorption of large proteins – such as Shh – when attached only to a single membrane anchor 3. Once desorbed, lipidated proteins can either reintegrate into the plasma membrane or associate with lipophilic soluble competitors, as shown in Fig. S7G. This demonstrates that another important role of the fully conserved dual N- and C-terminal Hh lipidation is to reduce non-enzymatic desorption of morphogens from the plasma membrane.

Earlier version: Cholesteroylated C-terminal peptides are required and sufficient for Disp-mediated protein transfer to HDL.

Revised version: Cholesterylated C-terminal peptides are necessary and sufficient for Disp-mediated protein transfer to HDL.

Earlier version: A) Immunoblotted cellular (c) and medium (m) fractions of ShhN-expressing nt ctrl and Disp-/- cells. Shown is unspecific ShhN solubilization into serum-free media (top blot, labeled [1]) or into serum-free DMEM supplemented with HDL (bottom blot, labeled [2]). A’) ShhN [1] expressed under serum-free conditions is solubilized in a monomeric state. A’’) ShhN [2] expressed in the presence of HDL remained monomeric (e.g. is not HDL associated). B) Cholesteroylated C25SShh solubilization into serum-free medium is strongly impaired (top blot, labeled [3]), but increases Dispdependently in the presence of HDL (bottom blot, labeled [4]). Asterisks denote C25SShh solubilized independent of Disp function. B’) Most C25SShh [3] in serum-free media is monomeric. B’’) C25SShh [4] expressed in the presence of HDL increases in molecular weight to match the molecular weight range of HDL (orange line). C) SEC of C25SShh solubilized from Disp-expressing cells (solid green line) or from Disp-/- cells (dotted line) indicates Disp-independent physical desorption and unregulated HDL association of the monolipidated protein. D) SEC of cholesteroylated mCherry solubilized from nt ctrl cells (solid lines) or from Disp-/- cells (dotted lines) under the same conditions. Dashed lines indicate proteins solubilized under serum-free conditions. Note that most mCherry associates with HDL in a Dispmediated manner. E) C25SShh (green line) disassembles from HDL in 50% ethanol (bright green line) or in 0.1% Triton X-100 (bright green dashed line). HDL (orange line) disassembly under the same conditions is confirmed by ApoA1 size shift toward the monomeric 32 kDa protein (bright orange line).

Revised version: A) Immunoblotted cellular (c) and medium (m) fractions of ShhN-expressing nt ctrl and Disp-/- cells. Shown is unspecific ShhN solubilization into serum-free media (upper blot, labeled [1]) or into serum-free DMEM supplemented with HDL (lower blot, labeled [2]). A’) ShhN [1] expressed under serum-free conditions is solubilized in a monomeric state. A’’) ShhN [2] expressed in the presence of HDL remained monomeric (i.e. not HDL-associated). B) Solubilization of cholesterylated C25SShh into serum-free medium is strongly impaired (upper blot, labeled [3]), but increases in a Disp-dependent manner in the presence of HDL (lower blot, labeled [4]). Asterisks denote C25SShh solubilized independently of Disp function. B’) Most C25SShh [3] in serum-free media is monomeric. B’’) C25SShh [4] expressed in the presence of HDL increases in molecular weight to match the molecular weight range of HDL (orange line). C) SEC of C25SShh solubilized from Disp-expressing cells (solid green line) or from Disp-/- cells (dotted line) shows Disp-independent physical desorption and unregulated HDL association of the monolipidated protein. D) SEC of cholesteroylated mCherry solubilized from nt ctrl cells (solid lines) or from Disp-/- cells (dotted lines) under the same conditions. Dashed lines indicate proteins solubilized under serum-free conditions. Note that most of the mCherry associates with HDL in a Disp-mediated manner. E) C25SShh (green line) dissociates from HDL in 50% ethanol (bright green line) or in 0.1% Triton X-100 (bright green dashed line). The disassembly of HDL (orange line) under the same conditions is confirmed by the size shift of ApoA1 towards the monomeric 32 kDa protein (light orange line).

Earlier version: Is the cholesteroylated peptide sufficient for Disp-mediated transfer to HDL, or does the globular Shh ectodomain participate in the transfer 47, 65? To answer this question, we replaced most of the C25SShh ectodomain with mCherry N-terminally flanked by the Shh secretion signal and C-terminally flanked by the Shh cholesterol transferase/autoprocessing domain. This strategy resulted in the secretion of cholesteroylated mCherry to the cell surface and its association with the outer leaflet of the plasma membrane (Fig. S7H). As observed previously for Shh (Fig. 5A), mCherry remained monomeric or formed small assemblies if released into serumdepleted media (Fig. 7D, dashed violet line), and HDL shifted mCherry elution into fractions that also contained ApoA1 (solid violet line). This result supports that HDL serves as a soluble acceptor for C-terminally cholesteroylated peptides. SEC also revealed that loss of Disp function decreased the relative amounts of HDL-associated material (Fig. 7D, solid violet line) and increased the relative amounts of solubilized monomeric protein (Fig. 7D, dotted violet line). Therefore, cholesteroylated peptides represent the critical determinant for Disp-mediated Shh transport to HDL. In contrast, Shh ectodomain interactions with Disp are not essential for the transfer. We supported physical C25SShh interactions with HDL by the addition of EtOH or Triton X-100, both of which disassembled soluble high-MW aggregates into monomers or small assemblies (Fig. 7E, bright green lines indicate C25AShh after treatment and the bright orange line ApoA1 after treatment).

Revised version: Is the cholesteroylated peptide sufficient for Disp-mediated transfer to HDL or is the globular Shh ectodomain involved in the transfer 47, 65? To answer this question, we replaced most of the C25SShh ectodomain with mCherry flanked N-terminally by the Shh secretion signal and C-terminally by the Shh cholesterol transferase/autoprocessing domain. This strategy resulted in the secretion of cholesteroylated mCherry to the cell surface and its association with the outer leaflet of the plasma membrane (Fig. S7H). As previously observed for Shh (Fig. 5A), mCherry remained monomeric or formed small aggregates when released into serum-depleted media (Fig. 7D, dashed violet line), and HDL shifted mCherry elution to fractions that also contained ApoA1 (solid violet line). This result supports that HDL serves as a soluble acceptor for C-terminally cholesterylated peptides. SEC also revealed that loss of Disp function decreased the relative amount of HDL-associated material (Fig. 7D, solid violet line) and increased the relative amount of solubilized monomeric protein (Fig. 7D, dotted violet line). Thus, cholesterylated peptides represent the critical determinant for Disp-mediated Shh transport to HDL. In contrast, Shh ectodomain interactions with Disp are not essential for the transfer. We supported the physical interactions of C25SShh with HDL by the addition of EtOH or Triton X-100, both of which dissociate soluble high MW aggregates into monomers or small assemblies (Fig. 7E, light green lines indicate C25AShh after treatment and the light orange line ApoA1 after treatment).

Discussion

Dual Shh lipidation and Disp-regulated Shh solubilization are interdependent

Earlier version: In the past, it was firmly established that Hh solubilization and signaling depend on Disp co-expression in source cells 5, 6, 66. However, the additional involvement of many cell surface and soluble co-factors has complicated efforts to unravel the exact nature of Disp-mediated Hh solubilization. To address this problem, we transfected Disp-deficient and Disp-expressing source cells 4 with Shh and systematically modified Shh solubilization by a number of extracellular co-factors. Cell-surface-associated Shh upstream of Disp function and soluble Shh downstream of Disp function were then analyzed together with their co-factors by using unbiased biochemical tools. The first important result of this approach is that both Shh lipids act together to prevent uncontrolled Shh desorption or shedding from producing cells. In contrast, artificially produced C25S/C25AShh that lacks N-palmitate and ShhN that lacks C-cholesterol are constitutively solubilized either by non-enzymatic desorption or by unregulated shedding (Fig. 8A). The second important result is that Disp and Scube2 specifically and synergistically increase shedding of the dual-lipidated precursor. These two observations suggest that cellsurface shedding constitutes a “ground state” from which dual-lipidated Shh is protected by tight plasma membrane association of both lipids – but only until Disp and Scube2 render Shh prone to proteolytic processing of the terminal lipidated peptides. This concept is important for the reinterpretation of the observation that dual lipidation is essential for unperturbed Hh biofunction in vivo 39, 53, 54, 55, 56. The current interpretation of these observations is that both Hh lipids contribute directly to Ptch1 receptor binding and maximal signaling. Our results support an alternative mechanism acting upstream of Ptch1: We suggest that the selective artificial prevention of N-or C-terminal lipidation during Hh/Shh biosynthesis in vivo (as achieved by the deletion of the N-terminal palmitate acceptor cysteine or of Hhat, or by preventing C-cholesteroylation 39, 53, 54, 55, 56) may have rendered spatiotemporally controlled Hh/Shh solubilization into unregulated “leakage” of monolipidated precursors from producing cells and tissues (Fig. 1). The “non-availability” of dual-lipidated Shh for Disp-controlled spatio-temporal signaling, or the desensitization of target cells to ongoing Shh exposure 67, may then have caused the observed Hh loss-of-function phenotypes 39, 53, 54, 55, 56, 68. For this reason, our release model does not contradict the established essential role of unimpaired dual Shh lipidation during biosynthesis for Hh activity regulation in vivo, but provides an alternative explanation for it. Published support for our model of Disp- and Scube2-regulated Hh shedding and release on HDL includes the following: Disp is required only in Hh-producing cells and not in receiving cells 64; Hh signaling defects in Disp-deficient model organisms are caused by a defect in the deployment of cell-surface associated Shh 5, 6, 7; and full Disp activity requires Scube2 10. Further support comes from the in vivo observation that artificially de-palmitoylated Hh variants impair development to variable degrees, but also show increased signaling in some tissues and induce ectopic signaling 40, 56. The latter observation is difficult to align with proposed essential signaling functions of the Shh palmitate, but is compatible with our model of spatiotemporally perturbed release of monolipidated proteins.

Revised version: In the past, it has been well established that Hh solubilization and signaling depend on Disp co-expression in source cells 5, 6, 66. However, the additional involvement of many cell surface and soluble cofactors has complicated efforts to elucidate the exact nature of Disp-mediated Hh solubilization. To address this problem, we transfected Disp-deficient and Disp-expressing source cells 4 with Shh and systematically modified Shh solubilization by a number of extracellular cofactors. Cell surface-associated Shh upstream of Disp function and soluble Shh downstream of Disp function were then analyzed together with their cofactors by using unbiased biochemical tools. The first important result of this approach is that both Shh lipids act together to prevent uncontrolled Shh desorption or shedding from producing cells. In contrast, engineered C25S/C25AShh, which lacks N-palmitate, and ShhN, which lacks C-cholesterol, are constitutively solubilized by either non-enzymatic desorption or unregulated shedding (Fig. 8A). The second important finding is that Disp and Scube2 specifically and synergistically increase the shedding of the dual-lipidated cellular protein. These two observations suggest that cell surface shedding represents a “ground state” from which dual-lipidated Shh is protected by tight plasma membrane association of both lipids – but only until Disp and Scube2 render Shh susceptible to proteolytic processing of the terminal lipidated peptides. This concept is important in reinterpreting the observation that dual lipidation is essential for unimpaired Hh biofunction in vivo 39, 53, 54, 55, 56. The current interpretation of these observations is that both Hh lipids contribute directly to Ptch1 receptor binding and maximal signaling. Our results support an alternative mechanism acting upstream of Ptch1: We suggest that the selective artificial prevention of N- or C-terminal lipidation during Hh/Shh biosynthesis in vivo (as achieved by the deletion of the N-terminal palmitate acceptor cysteine or of Hhat, or by prevention of C-cholesteroylation 39, 53, 54, 55, 56) may have converted spatiotemporally controlled Hh/Shh solubilization into unregulated “leakage” of monolipidated precursors from producing cells and tissues (Fig. 1). The “unavailability” of dual-lipidated Shh for Disp-controlled spatiotemporal signaling, or the desensitization of target cells to ongoing Shh exposure 67, may then have caused the observed Hh loss-of-function phenotypes 39, 53, 54, 55, 56, 68. Therefore, our release model does not contradict the established essential role of unimpaired dual Shh lipidation during biosynthesis for the regulation of Hh activity in vivo, but provides an alternative explanation for it. Published support for our model of Disp- and Scube2-regulated Hh shedding and release on HDL includes the following: Disp is required only in Hh-producing cells and not in receiving cells 64; Hh signaling defects in Disp-deficient model organisms are caused by a defect in the deployment of cell surface-associated Shh 5, 6, 7; and full Disp activity requires Scube2 10. Further support comes from the in vivo observation that artificially depalmitoylated Hh variants impair development to varying degrees, but also show increased signaling in some tissues and induce ectopic signaling 40, 56. The latter observation is difficult to reconcile with the proposed essential signaling functions of the Shh palmitate, but is compatible with our model of spatiotemporally perturbed release of monolipidated proteins.

Model of two-way Disp-mediated Shh solubilization.

Earlier version: A) Dual lipidation protects Shh from unregulated cell-surface shedding by tight plasma membrane association of both lipids (blocked shedding is indicated by an x in [1]). In contrast, monolipidated ShhN [2] and C25SShh [3] are prone to unregulated membrane-proximal shedding (indicated by the dashed line) or non-enzymatic desorption. B) Representation of surface hydrophobicity of Disp (pdb:7RPH 47) suggests an extended hydrophobic surface channel (hydrophobic residues are shown in red) that may function as a “slide” for lipophiles extending from the plasma membrane (dashed lines) to a cavity of the second extracellular domain (blue arrows). An upward lifted sterol (green stick representation) at the start point of the hydrophobic track may represent an intermediate state of sterol extraction from the membrane, and a lipidic group modeled as the amphiphilic detergent lauryl maltose neopentyl glycol (violet stick structure) may represent the end point of the transfer 47 prior to handover to HDL. C) We suggest two sequences of Shh transfer events. In the first event [1], plasma membrane sterol is transferred through the Disp hydrophobic surface channel to HDL acceptors. This process resembles established reverse cholesterol transport. In the second event, if present, C-terminal cholesterol moieties of Shh can also be transferred [2]. This partial Shh extraction exposes the N-terminal cleavage site [3] and makes it prone to proteolytic processing (similar to ShhN as shown in A). N-terminal Shh shedding may then release the protein from the plasma membrane [4] to complete the transfer [5]. In addition to, or competing with this process, cholesterol depletion of the plasma membrane (representing the first event, [1]) may trigger shedding of both terminal Shh peptides and the solubilization of monomeric proteins in an indirect manner [6]. See Discussion for details.

Revised version: A) Dual lipidation protects Shh from unregulated cell surface shedding by tight plasma membrane association of both lipids (blocked shedding is indicated by an x in [1]). In contrast, monolipidated ShhN [2] and C25SShh [3] are prone to unregulated membrane proximal shedding (indicated by the dashed line) or non-enzymatic desorption. B) Depiction of the surface hydrophobicity of Disp (pdb:7RPH 47) suggests an extended hydrophobic surface channel (hydrophobic residues are shown in red) that may function as a “slide” for lipophiles extending from the plasma membrane (dashed lines) to a cavity of the second extracellular domain (blue arrows). A sterol lifted upward (green stick representation) at the starting point of the hydrophobic track may represent an intermediate state of sterol extraction from the membrane, and a lipidic group modeled as the amphiphilic detergent lauryl maltose neopentyl glycol (violet stick structure) may represent the end point of the transfer 47 prior to delivery to HDL. C) We propose two sequences of Shh transfer events. In the first event [1], plasma membrane sterol is transferred through the hydrophobic Disp surface channel to HDL acceptors. This process is similar to the established reverse cholesterol transport. In the second event, if present, C-terminal cholesterol moieties of Shh can also be transferred [2]. This partial Shh extraction exposes the N-terminal cleavage site [3] and makes it susceptible to proteolytic processing (similar to ShhN as shown in A). N-terminal Shh shedding can then release the protein from the plasma membrane [4] to complete the transfer [5]. In addition to, or competing with this process, cholesterol depletion of the plasma membrane (representing the first event, [1]) may indirectly trigger shedding of both terminal Shh peptides and the solubilization of monomeric proteins [6]. See Discussion for details.

Earlier version: Disp mediates Shh solubilization by two functionally related transfer modes

Revised version: Disp mediates Shh solubilization through two functionally related modes of transfer

Earlier version: The second important result of our study is that the absence of all serum traces reduces or even abolishes Disp/Scube2-mediated Shh solubilization. This reveals that Disp and Scube2 are insufficient to solubilize dual-lipidated Shh and that a serum factor is also required. We showed previously that Disp knockout in cells impairs 3Hcholesterol efflux and increases cholesterol levels in membranes, and that this in turn inhibits Shh shedding in an indirect manner 4. We also showed that HDL serves as a soluble sink for Disp-exported cholesterol and that Shh shedding from Disp-/- cells was restored by Disp and the established cholesterol transporter Ptch1 4.

Revised version: The third important finding of our study is that the absence of all serum traces reduces or even abolishes Disp/Scube2-mediated Shh solubilization. This indicates that Disp and Scube2 are not sufficient to solubilize dual-lipidated Shh and that a serum factor is also required. We have previously shown that Disp knockout in cells impairs 3Hcholesterol efflux and increases membrane cholesterol levels, which in turn indirectly inhibits Shh shedding 4. We also showed that HDL serves as a soluble sink for Disp-exported cholesterol and that Shh shedding from Disp-/- cells is restored by Disp and the established cholesterol transporter Ptch1 4. Consistent with these published observations, both Disp and Ptch1 contain sterol-sensing domains that are likely involved in cholesterol transport or its regulation (Fig. 8B). Taken together, these published findings suggest that the required serum factor is HDL.

Removed in the revised version: Consistent with these published observations, both Disp and Ptch1 contain sterol-sensing domains likely involved in cholesterol transport or in its regulation. Together, these published findings suggest that the required serum factor is HDL.

Removed in the revised version: Here we describe that HDL also accepts cholesteroylated – but not dual-lipidated – Shh from Disp (Fig. 8C). This finding is supported by the presence of an extended hydrophobic surface channel in Disp that may function as an open “slide” for lipophiles (Fig. 8B). This slide may be powered by a transmembrane Na+ flux 47, similar to the H+ flux that drives related prokaryotic resistance-nodulation-division transporter export of small molecules 69. Disp exporter function of cholesteroylated proteins to HDL is further supported by the published concept that the fly LPP lipophorin carries cholesteroylated Hh in vivo 13, 14, 15. We extend this published concept by showing that the most C-terminal cholesteroylated Shh peptide is sufficient for direct Disp-mediated protein export to HDL because mCherry linked to this peptide also transfers to HDL in a mostly Disp-dependent manner. This is in line with a previous report demonstrating that C-cholesterol is required and sufficient for Disp-mediated protein export 10.

Earlier version: Because of their small size of 5-10 nm, HDLs are not only abundant in the circulation, but are also present in interstitial fluids, both in the adult and throughout development 14. In contrast, larger LPPs such as LDLs and VLDLs are restricted in their distribution by vascular impermeability 52, 70. The ratio of HDL to LDL in interstitial fluids has been estimated to average 50:1 70. These properties make HDL well suited not only for reverse cholesterol transport, but also for the transport of cholesteroylated Hh/Shh cargo to its receptor Ptch1 that also acts as an LPP receptor 71. This would lead to the interesting concept that “Hh-free” LPP and “Hh-loaded” LPPs may compete for Ptch1 receptor binding, the former by promoting Ptch1-mediated cholesterol export to suppress signaling and the latter by terminating cholesterol efflux to trigger Hh signaling 43. Indeed, this elegant mechanism of Ptch1 activity regulation has previously been demonstrated, both in vitro and in vivo 14.

Revised version: In this paper, we describe that HDL also accepts cholesteroylated – but not dual-lipidated – Shh from Disp (Fig. 8C). This finding is supported by the presence of an extended hydrophobic surface channel in Disp that may function as an open “slide” for lipophiles (Fig. 8B, blue arrows). This slide may be powered by a transmembrane Na+ flux 47, similar to the H+ flux that drives the related prokaryotic resistance-nodulation-division transporter export of small molecules 69. The Disp exporter function of cholesteroylated proteins to HDL is further supported by the published concept that the fly LPP lipophorin carries cholesteroylated Hh in vivo 13, 14, 15. We extend this published concept by showing that the most C-terminal cholesteroylated Shh peptide is sufficient for direct Disp-mediated protein export to HDL, because mCherry linked to this peptide is also transferred to HDL in a mostly Disp-dependent manner. This is consistent with a previous report showing that C-cholesterol is necessary and sufficient for Disp-mediated protein export 10. Due to their small size of 5-10 nm, HDLs are not only abundant in the circulation, but are also present in interstitial fluids, both in the adult and during development 14. In contrast, larger LPPs such as LDLs and VLDLs are limited in their distribution by vascular impermeability 52, 70. The ratio of HDL to LDL in interstitial fluids has been estimated to average 50:1 70. These properties make HDL well suited not only for reverse cholesterol transport, but also for the transport of cholesteroylated Hh/Shh cargo to its receptor Ptch1, which also functions as an LPP receptor 71. This would lead to the interesting concept that “Hh-free” LPPs and “Hh-loaded” LPPs may compete for Ptch1 receptor binding, the former by promoting Ptch1-mediated cholesterol export to suppress signaling and the latter by terminating cholesterol efflux to trigger Hh signaling 43. Indeed, this elegant mechanism of Ptch1 activity regulation has been previously demonstrated both in vitro and in vivo 14.

Earlier version: N-terminal shedding during Disp-mediated Hh export in vivo

Revised version: Shedding of the N-terminus during Disp-mediated export of Hh in vivo

Earlier version: Our in vitro-derived concept of required N-terminal peptide shedding is supported by the in vivo finding that site-directed mutagenesis of the sheddase target site fully abolishes transgene function in the developing Drosophila wing and eye (see Fig. S8 for a detailed description of repeated and combined experiments from published work 34, 35). Another striking observation was that the same mutated proteins suppress endogenous Hh biofunction in a dominant-negative manner 34, 35. The same site-directed mutagenesis approach on the Hh C-terminus, however, affects biofunction of the transgene only mildly and does not suppress endogenous Hh biofunction in vivo. Both observations can now be explained by mutated Hh transgenes having their C-terminal peptides readily associated with the extended Disp “slide” for lipophiles, with the N-mutated protein unable to complete the transfer due to blocked N-terminal shedding and continued plasma membrane association of the palmitoylated Hh N-terminus (preventing step [3] and [4] because of transfer arrest at step [2] in Fig. 8C). As a consequence, the resulting Disp bottleneck would slow down endogenous Hh release, explaining the observed dominant-negative developmental defects 34, 35 (Fig. S8C,D). This sequence of events is supported by reversed dominant-negative defects of the N-mutant protein upon additional removal of the palmitate membrane anchor 34, 35. Therefore, taken together, our findings suggest that most Hh solubilization in vivo does not require C-terminal Hh shedding, but requires direct cholesteroylated Hh handover from Disp to LPPs. During this process, only the palmitoylated N-terminus is shed to complete the transfer (Fig. 8C, Fig. S8). Our finding that cholesteroylated C-terminal peptides are sufficient for Disp-mediated transfer to HDL is supported by the in vivo observation that transgenes with cholesteroylated 27 kDa green fluorescent protein tags downstream of the 19 kDa Hh signaling domain are bioactive in flies 72 and mice 73. Yet, Disp specificity in vivo can be elegantly explained by Hh being the only known metazoan protein with covalently linked cholesterol. Finally, our finding that palmitoylated Shh N-termini are not extracted and relayed by Disp is supported by the in vivo observation that transgenic expression of ShhN – the variant protein that is N-palmitoylated but lacks C-cholesterol – rescues many early Hh-related embryonic lethal defects in Disp-/- mutant mice 64.

Revised version: Our in vitro-derived concept of required N-terminal peptide shedding is supported by the in vivo finding that site-directed mutagenesis of the sheddase target site completely abolishes transgene function in the developing Drosophila wing and eye (see Fig. S8 for a detailed description of repeated and combined experiments from published work 34, 35). Another striking observation was that the same mutant proteins suppress endogenous Hh biofunction in a dominant-negative manner 34, 35. However, the same site-directed mutagenesis approach at the Hh C-terminus only slightly affects the biofunction of the transgene and does not suppress endogenous Hh biofunction in vivo. Both observations can now be explained by the fact that mutant Hh transgenes readily associate their C-terminal peptides with the extended Disp “slide” for lipophiles, with the N-mutant protein unable to complete the transfer due to blocked N-terminal shedding and continued plasma membrane association of the palmitoylated Hh N-terminus (preventing step [3] and [4] due to transfer arrest at step [2] in Fig. 8C). As a consequence, the resulting Disp bottleneck would slow down endogenous Hh release, explaining the observed dominant-negative developmental defects 34, 35 (Fig. S8C,D). This sequence of events is supported by reversed dominant-negative defects of the N mutant protein upon additional removal of the palmitate membrane anchor 34, 35. Taken together, our results suggest that most Hh solubilization in vivo does not require C-terminal Hh shedding, but rather direct cholesterylated Hh transfer from Disp to LPPs. During this process, only the palmitoylated N-terminus is shed to complete the transfer (Fig. 8C, Fig. S8). Our finding that cholesterylated C-terminal peptides are sufficient for Disp-mediated transfer to HDL is supported by the in vivo observation that transgenes with cholesterylated 27 kDa green fluorescent protein tags downstream of the 19 kDa Hh signaling domain are bioactive in flies 72 and mice 73. However, the Disp specificity in vivo can be elegantly explained by the fact that Hh is the only known metazoan protein with covalently attached cholesterol. Finally, our finding that palmitoylated Shh N-termini are not extracted and translocated by Disp is supported by the in vivo observation that transgenic expression of ShhN – the variant protein that is N-palmitoylated but lacks C-cholesterol – rescues many early Hh-related embryonic lethal defects in Disp-/- mutant mice 64.

Supporting information